The Rhizobial effector NopT targets Nod factor receptors to regulate symbiosis in Lotus japonicus

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #1 (Public review):

Bacterial effectors that interfere with the inner molecular workings of eukaryotic host cells are of great biological significance across disciplines. On the one hand they help us to understand the molecular strategies that bacteria use to manipulate host cells. On the other hand they can be used as research tools to reveal molecular details of the intricate workings of the host machinery that is relevant for the interaction/defence/symbiosis with bacteria. The authors investigate the function and biological impact of a rhizobial effector that interacts with and modifies, and curiously is modified by, legume receptors essential for symbiosis. The molecular analysis revealed a bacterial effector that cleaves a plant symbiosis signaling receptor to inhibit signaling and the host counterplay by phosphorylation via a receptor kinase. These findings have potential implications beyond bacterial interactions with plants.

Bao and colleagues investigated how rhizobial effector proteins can regulate the legume root nodule symbiosis. A rhizobial effector is described to directly modify symbiosis-related signaling proteins, altering the outcome of the symbiosis. Overall, the paper presents findings that will have a wide appeal beyond its primary field.

Out of 15 identified effectors from Sinorhizobium fredii, they focus on the effector NopT, which exhibits proteolytic activity and may therefore cleave specific target proteins of the host plant. They focus on two Nod factor receptors of the legume Lotus japonicus, NFR1 and NFR5, both of which were previously found to be essential for the perception of rhizobial nod factor, and the induction of symbiotic responses such as bacterial infection thread formation in root hairs and root nodule development (Madsen et al., 2003, Nature; Tirichine et al., 2003; Nature). The authors present evidence for an interaction of NopT with NFR1 and NFR5. The paper aims to characterize the biochemical and functional consequences of these interactions and the phenotype that arises when the effector is mutated.

Evidence is presented that in vitro NopT can cleave NFR5 at its juxtamembrane region. NFR5 appears also to be cleaved in vivo. and NFR1 appears to inhibit the proteolytic activity of NopT by phosphorylating NopT. When NFR5 and NFR1 are ectopically over-expressed in leaves of the non-legume Nicotiana benthamiana, they induce cell death (Madsen et al., 2011, Plant Journal). Bao et al., found that this cell death response is inhibited by the coexpression of nopT. Mutation of nopT alters the outcome of rhizobial infection in L. japonicus. These conclusions are well supported by the data.

The authors present evidence supporting the interaction of NopT with NFR1 and NFR5. In particular, there is solid support for cleavage of NFR5 by NopT (Figure 3) and the identification of NopT phosphorylation sites that inhibit its proteolytic activity (Figure 4C). Cleavage of NFR5 upon expression in N. benthamiana (Figure 3A) requires appropriate controls (inactive mutant versions) that have been provided, since Agrobacterium as a closely rhizobia-related bacterium might increase defense related proteolytic activity in the plant host cells.

We appreciate your recognition of the importance of appropriate controls in our experimental design. In response to your comments, we revised our manuscript to ensure that the figures and legends provide a clear description of the controls used. We also included a more detailed description of our experimental design at several places. In particular, we have highlighted the use of the protease-dead version of NopT as a control (NopTC93S). Therefore, NFR5-GFP cleavage in N. benthamiana clearly depended on protease activity of NopT and not on Agrobacterium (Fig. 3A). In the revised text, we carefully revied the conclusion and do not conclude at this stage that NopT proteolyzes NFR5. However, our subsequent experiments, including in vitro experiments, clearly show that NopT is able to proteolyze NFR5.

Key results from N. benthamiana appear consistent with data from recombinant protein expression in bacteria. For the analysis in the host legume L. japonicus transgenic hairy roots were included. To demonstrate that the cleavage of NFR5 occurs during the interaction in plant cells the authors build largely on western blots. Regardless of whether Nicotiana leaf cells or Lotus root cells are used as the test platform, the Western blots indicate that only a small proportion of NFR5 is cleaved when co-expressed with nopT, and most of the NFR5 persists in its full-length form (Figures 3A-D). It is not quite clear how the authors explain the loss of NFR5 function (loss of cell death, impact on symbiosis), as a vast excess of the tested target remains intact. It is also not clear why a large proportion of NFR5 is unaffected by the proteolytic activity of NopT. This is particularly interesting in Nicotiana in the absence of Nod factor that could trigger NFR1 kinase activity.

Thank you for your comments regarding the cleavage of NFR5 by NopT and its functional implications. We acknowledge that our immunoblots indicate only a relatively small proportion of the NFR5 cleavage product. Possible explanations could be as follows:

(1) The presence of full-length NFR5 does not preclude a significant impact of NopT on function of NFR5, as NopT is able to interact with NFR5. In other words, the NopT-NFR5 and NopT-NFR1 interactions at the plasma membrane might influence the function of the NFR1/NFR5 receptor without proteolytic cleavage of NFR5. In fact, protease-dead NopTC93S expressed in NGR234ΔnopT showed certain effects in L. japonicus (less infection foci were formed compared to NGR234ΔnopT Fig. 5E). In this context, it is worth mentioning that the non-acylated NopTC93S (Fig. 1B) and NopTUSDA257 (Fig. 6B) proteins were unable to suppress NFR1/NFR5-induced cell death in N. benthamina, but this could be explained by the lack of acylation and altered subcellular localization.

(2) In the cleavage assay, only small portion of NFR5 could be detected for cleavage by NopT. However, this cleavage might be sufficient to suppress signaling pathways, leading to the observed phenotypic changes (loss of cell death in N. benthamiana; altered infection in L. japonicus). We do believe this is a great point, therefore, we carefully revised the conclusion about this point. Throughout the paper, we stated that the cleavage of NFR5 suppresses symbiotic signaling but not disrupt the symbiotic signaling. We also removed the conclusion that cleavage of NFR5 by NopT results in the function loss of NFR5.

(3) N. benthamiana co-expressing NFR1/NFR5 leads to strong cell death, which suggest that the NFR1 kinase activity might be constitutively active even in the absence of Nod factors. But why co-expression of symbiotic receptor leads to cell death and how kinase activity is active in the absence of Nod factor are not clear, which is of great interest to be studied.

(4) The proteolytic activity of NopT may be reduced by the interaction of NopT with other proteins such as NFR1, which phosphorylates NopT and inactivates its protease activity.

In our revised manuscript version, we provide now quantitative data for the efficiency of NFR5 cleavage by NopT in different expression systems used (Figure 3 and Supplemental Fig. 16). We have also improved our Discussion in this context.

Comments on latest version:

The presentation of the figures and the language has greatly improved and the specific mistakes pointed out in the last review have been corrected. I especially appreciate the new images used to illustrate the observed mutant phenotypes, which are much clearer and easier to understand. The pictures used to illustrate the mutant phenotypes seem to be of more comparable root regions than before. Overall, the requested changes have been implemented, with some exceptions described below.

• Figure 1: New representative images are shown for BAX1 and CERK1. These pictures are more consistent with the phenotype seen in other treatments, but since the data has not changed, I presume the data from leaf discs (where the leaf discs for these treatments looked very different) previously shown is still included. The criteria for what was considered cell death is in my opinion still not described in the legend. The cell death/total ratio has been added for all leaf discs, as requested.

Thank you so much for carefully pointing out this. Cell death in leaf disc results in the formation of necrotic plaques, which restrains pathogens within deceased cells. These plaques commonly manifest as leaf dehydration, frequently accompanied by a translucent appearance. Brown and shriveled leaf discs serve as indicators of cell death. We have added these descriptions in the figure legend of Figure 1.

• Figure 2: the discussion of the figure now emphasizes direct protein interaction. There is still no size marker in 2D or a description of size in the figure legend, making it difficult to compare the result to Figure 3. If I understand the rebuttal comments correctly, there are other bands on the blot, including non-specific bands. This does not negate the need to include the full blot as a supplemental figure to show cleaved NFR5 as well as other bands. I do not see any other clarifications on this subject in the manuscript.

Thank you for your suggestion. In the revised manuscript, we have included the kDa range for all proteins detected in Figure.2D. The full blot of Co-IP assay was shown in Fig S2 (a new supplemental data). Yes, we detected some smaller bands after immunoblot, but we cannot give clear conclusion of what these bands are based on the current study. Interestingly, these smaller bands were immunoprecipitated by anti-FLAG beads, suggesting that these bands are some truncated peptides from NFR5.

• Figure 5: From the pictures, it is now easier to understand what is meant by "infection foci". Although there is no description in the methods of how these were distinguished from infection threads, I believe the images are clear enough.

Thank you for your helpful comment. In the revised manuscript, we have added the descriptions about this experiment in the method section and in the legend in Figure 5A.

• Figure 6: The changes in the discussion are appreciated, but panel E still misrepresents the evidence in the paper, as from the drawing it still seems that the cleaved NFR5 is somehow directly responsible for suppressing infection when this was not shown.

Thank you for your thoughtful comments. We appreciate your suggestion to the schematic model to illustrate the cleavage of NFR5 to suppressing rhizobia infection. In the revised manuscript, we have changed the model in Figure 6E.

Reviewer #2 (Public review):

Summary:

This manuscript presents data demonstrating NopT's interaction with Nod Factor Receptors NFR1 and NFR5 and its impact on cell death inhibition and rhizobial infection. The identification of a truncated NopT variant in certain Sinorhizobium species adds an interesting dimension to the study. These data try to bridge the gaps between classical Nod-factor-dependent nodulation and T3SS NopT effector-dependent nodulation in legume-rhizobium symbiosis. Overall, the research provides interesting insights into the molecular mechanisms underlying symbiotic interactions between rhizobia and legumes.

Strengths:

The manuscript nicely demonstrates NopT's proteolytic cleavage of NFR5, regulated by NFR1 phosphorylation, promoting rhizobial infection in L. japonicus. Intriguingly, authors also identify a truncated NopT variant in certain Sinorhizobium species, maintaining NFR5 cleavage but lacking NFR1 interaction. These findings bridge the T3SS effector with the classical Nod-factor-dependent nodulation pathway, offering novel insights into symbiotic interactions.

Weaknesses:

(1) In the previous study, when transiently expressed NopT alone in Nicotiana tobacco plants, proteolytically active NopT elicited a rapid hypersensitive reaction. However, this phenotype was not observed when expressing the same NopT in Nicotiana benthamiana (Figure 1A). Conversely, cell death and a hypersensitive reaction were observed in Figure S8. This raises questions about the suitability of the exogenous expression system for studying NopT proteolysis specificity.

We appreciate your attention to these plant-specific differences. Previous studies showed that NopT expressed in tobacco (N. tabacum) or in specific Arabidopsis ecotypes (with PBS1/RPS5 genes) causes rapid cell death (Dai et al. 2008; Khan et al. 2022). Khan et al. 2022 reported recently that cell death does not occur in N. benthamiana unless the leaves were transformed with PBS1/RPS5 constructs. Our data shown in Fig. S17 confirm these findings. As cell death is usually associated with induction of plant protease activities, we considered N. tabacum and A. thaliana plants as not suitable for testing NFR5 cleavage by NopT. In fact, no NopT/NFR5 experiments were not performed with these plants in our study. In response to your comment, we now better describe the N. benthamiana expression system and cite the previous articles_._ Furthermore, we have revised the Discussion section to better emphasize effector-induced immunity in non-host plants and the negative effect of rhizobial effectors during symbiosis. Our revisions certainly provide a clearer understanding of the advantages and limitations of the N. benthamiana expression system.

(2) NFR5 Loss-of-function mutants do not produce nodules in the presence of rhizobia in lotus roots, and overexpression of NFR1 and NFR5 produces spontaneous nodules. In this regard, if the direct proteolysis target of NopT is NFR5, one could expect the NGR234's infection will not be very successful because of the Native NopT's specific proteolysis function of NFR5 and NFR1. Conversely, in Figure 5, authors observed the different results.

Thank you for this comment, which points out that we did not address this aspect precisely enough in the original manuscript version. We improved our manuscript and now write that nfr1 and nfr5 mutants do not produce nodules (Madsen et al., 2003; Radutoiu et al., 2003) and that over-expression of either NFR1 or NFR5 can activate NF signaling, resulting in formation of spontaneous nodules in the absence of rhizobia (Ried et al., 2014). In fact, compared to the nopT knockout mutant NGR234ΔnopT, wildtype NGR234 (with NopT) is less successful in inducing infection foci in root hairs of L. japonicus (Fig. 5). With respect to formation of nodule primordia, we repeated our inoculation experiments with NGR234ΔnopT and wildtype NGR234 and also included a nopT over-expressing NGR234 strain into the analysis. Our data clearly showed that nodule primordium formation was negatively affected by NopT. The new data are shown in Fig. 5 of our revised version. Our data show that NGR234 infection is not really successful, especially when NopT is over-expressed. This is consistent with our observations that NopT targets Nod factor receptors in L. japonicus and inhibits NF signaling (NIN promoter-GUS experiments). Our findings indicate that NopT might be an “Avr effector” for L. japonicus. However, in other host plants of NGR234, NopT possesses a symbiosis-promoting role (Dai et al. 2008; Kambara et al. 2009). Such differences could be explained by different NopT targets in different plants (in addition to Nod factor receptors), which may influence the outcome of the infection process. Indeed, our work shows that NopT can interact with various kinase-dead LysM domain receptors, suggesting a role of NopT in suppression or activation of plant immunity responses depending on the host plant. We discuss such alternative mechanisms in our revised manuscript version and emphasize the need for further investigation to elucidate the precise mechanisms underlying the observed infection phenotype and the role of NopT in modulating symbiotic signaling pathways. In this context, we would also like to mention the new figures of our manuscript which are showing (i) the efficiency of NFR5 cleavage by NopT in different expression systems (Figure 3), (ii) the interaction between NopTC93S and His-SUMO-NFR5JM-GFP (Supplementary Fig. 5), and (iii) cleavage of His-SUMO-NFPJM-GFP by NopT (Supplementary Figs. S8 and S9).

(3) In Figure 6E, the model illustrates how NopT digests NFR5 to regulate rhizobia infection. However, it raises the question of whether it is reasonable for NGR234 to produce an effector that restricts its own colonization in host plants.

Thank you for mentioning this point. We are aware of the possible paradox that the broad-host-range strain NGR234 produces an effector that appears to restrict its infection of host plants. As mentioned in our answer to the previous comment, NopT could have additional functions beyond the regulation of Nod factor signaling. In our revised manuscript version, we have modified our text as follows:

(1) We mention the potential evolutionary aspects of NopT-mediated regulation of rhizobial infection and discuss the possibility that interactions between NopT and Nod factor receptors may have evolved to fine-tune Nod factor signaling to avoid rhizobial hyperinfection in certain host legumes.

(2) We also emphasize that the presence of NopT may confer selective advantages in other host plants than L. japonicus due to interactions with proteins related to plant immunity. Like other effectors, NopT could suppress activation of immune responses (suppression of PTI) or cause effector-triggered immunity (ETI) responses, thereby modulating rhizobial infection and nodule formation. Interactions between NopT and proteins related to the plant immune system may represent an important evolutionary driving force for host-specific nodulation and explain why the presence of NopT in NGR234 has a negative effect on symbiosis with L. japonicus but a positive one with other legumes.

(4) The failure to generate stable transgenic plants expressing NopT in Lotus japonicus is surprising, considering the manuscript's claim that NopT specifically proteolyzes NFR5, a major player in the response to nodule symbiosis, without being essential for plant development.

We also thank for this comment. We have revised the Discussion section of our manuscript and discuss now our failure to generate stable transgenic L. japonicus plants expressing NopT. We observed that the protease activity of NopT in aerial parts of L. japonicus had a negative effect on plant development, whereas NopT expression in hairy roots was possible. Such differences may be explained by different NopT substrates in roots and aerial parts of the plant. In this context, we also discuss our finding that NopT not only cleaves NFR5 but is also able to proteolyze other proteins of L. japonicus such as LjLYS11, suggesting that NopT not only suppresses Nod factor signaling, but may also interfere with signal transduction pathways related to plant immunity. We speculate that, depending on the host legume species, NopT could suppress PTI or induce ETI, thereby modulating rhizobial infection and nodule formation.

Comments on revised version:

This version has effectively addressed most of my concerns. However, one key issue remains unresolved regarding the mechanism of NopT in regulating nodule symbiosis. Specifically, the explanation of how NopT catabolizes NFR5 to regulate symbiosis is still not convincing within the current framework of plant-microbe interaction, where plants are understood to genetically control rhizobial colonization.

While alternative regulatory mechanisms in plant-microbe interactions are plausible, the notion that the NRG234-secreted effector NopT could reduce its own infection by either suppressing plant immunity or degrading the symbiosis receptor remains unsubstantiated. I believe further revisions are needed in the discussion section to more clearly address and clarify these findings and any lingering uncertainties.

We appreciate your positive comments on the reason why NopT catabolizes NFR5 to regulate symbiosis. NopT belongs to pathogen effecftors YopT family and also cleavage Arabidopsis AtLYK5 and L. japonicus LjLYS11 which trigger immunity responses in plants. NFR5, AtLYK5 and LjLYS11 has the conserved amino acid motif at the juxtamembrane domain, leading to cleaving NFR5 by NopT during symbiosis. Besides, in plant-microbe interaction, effector HopB1 cleaves immune co-receptor BAK1 at the kinase domain to inhibit plant defense. The effect on cleavage of receptor may be positive or negative. NopT suppressing symbiosis may avoid preventing hyperinfection in the specific interaction between rhizobia and legumes. In the revised manuscript, we have emphasized this point more clearly in why NopT could reduce its own infection by either suppressing plant immunity in discussion.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

Evaluation of the author's responses to the reviewer comments during the first review round

Reviewer's Comment:

Regardless of whether Nicotiana leaf cells or Lotus root cells are used as the test platform, the Western blots indicate that only a small proportion of NFR5 is cleaved when co-expressed with NopT, and most of the NFR5 persists in its full-length form (Figures 3A-D). It is not quite clear how the authors explain the loss of NFR5 function (loss of cell death, impact on symbiosis), as a vast excess of the tested target remains intact. It is also not clear why a large proportion of NFR5 is unaffected by the proteolytic activity of NopT. This is particularly interesting in Nicotiana in the absence of Nod factor that could trigger NFR1 kinase activity.

Summary of response:

• NopT could be interfering with the NFR1/NFR5 complex without proteolytic cleavage

• The cleaved fraction may still be sufficient to disrupt signaling pathways

• Elevated abundance of NFR5 relative to WT levels

• Add quantitative data for efficiency of NFR5 cleavage in different systems

Evaluation of response:

• The quantification of NFR5 cleavage efficiency is welcome, and there is some discussion of the possible reasons for the large proportion of uncleaved NFR5. It is clear that there is a large difference in cleavage efficiency between L. japonicus roots and N. benthamiana.

• The data is shown as a bar plot. Given that only 3 biological replicates are used, the data points should be shown, and there is too little data to provide sensible error bars. It would be better to simply make a dot-plot and indicate the mean for each sample. However, the main aim of the comment is addressed.

Thank you for your constructive comments regarding Figure S16. In the revised manuscript, we have presented these data into dot-Plot format.

Reviewer's Comment:

It is also difficult to evaluate how the ratios of cleaved and full-length protein change when different versions of NopT are present without a quantification of band strengths normalized to loading controls (Figure 3C, 3D, 3F). The same is true for the blots supporting NFR1 phosphorylation of NopT (Figure 4A).

Summary of response:

• Quantified proportion of cleaved and full length NFR5 in different systems (S14)

• Band strengths of immunoblots quantified (4B)

Evaluation of response:

• The quantification has been performed as requested and the data is shown as bar plots. This type of data is frequently displayed as part of the blot figure itself, printed under each respective lane, making it easier for the reader to connect the ratios to the band sizes. If data is shown in a plot, the data points should be shown on the plot, as described above.

Thank you for your constructive comments regarding Figure 3. In the revised manuscript, we have added the cleavage efficiency in the 3A-3D.

Reviewer's Comment:

Nodule primordia and infection threads are still formed when L. japonicus plants are inoculated with ∆nopT mutant bacteria, but it is not clear if these primordia are infected or develop into fully functional nodules (Figure 5). A quantification of the ratio of infected and non-infected nodules and primordia would reveal whether NopT is only active at the transition from infection focus to thread or perhaps also later in the bacterial infection process of the developing root nodule.

Summary of response:

• Additional experiments with NGR234 or NGR234ΔnopT mutants find no non-infected nodules (fig. 5)

Evaluation of response:

• The requested quantification has been done, although the support for the findings would be stronger if also mature nodules per plant were quantified and plotted. If non-infected nodules were neither present in NGR234 or NGR234ΔnopT, it would still be advisable to include images of cross-sections of the fully-developed nodules.

We appreciate your positive comments on the cross-sections of the fully-developed nodules. In the revised manuscript, we have added the cross-section images of nodules in the Figure S12.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: It is well-documented that type-III effectors are required by Gram-negative pathogens to directly target different host cellular pathways to promote bacterial infection. However, in the context of legume-rhizobium symbiosis, the role of rhizobial effectors in regulating plant symbiotic pathways remains largely unexplored. Here, we show that NopT, a YopT-type cysteine protease of Sinorhizobium fredii NGR234 directly targets the plant’s symbiotic signaling pathway by associating with two Nod factor receptors (NFR1 and NFR5 of Lotus japonicus). NopT inhibits cell death triggered by co-expression of NFR1/NFR5 in Nicotiana benthamiana. Full-length NopT physically interacts with NFR1 and NFR5. NopT proteolytically cleaves NFR5 both in vitro and in vivo, but can be inactivated by NFR1 as a result of phosphorylation. NopT plays an essential role in mediating rhizobial infection in L. japonicus. Autocleaved NopT retains the ability to cleave NFR5 but no longer interacts with NFR1. Interestingly, genomes of certain Sinorhizobium species only harbor nopT genes encoding truncated proteins without the autocleavage site. These results reveal an intricate interplay between rhizobia and legumes, in which a rhizobial effector protease targets NFR5 to dampen symbiotic signaling. NFR1 appears to counteract this process by phosphorylating the effector. This discovery highlights the role of a bacterial effector in regulating a signaling pathway in plants and opens up the perspective of developing kinase-interacting proteases to fine-tune cellular signaling processes in general.

Revised version: It is well-documented that type-III effectors are required by Gram-negative pathogens to directly target different host cellular pathways to promote bacterial infection. However, in the context of legume-rhizobium symbiosis, the role of rhizobial effectors in regulating plant symbiotic pathways remains largely unexplored. Here, we show that NopT, a YopT-type cysteine protease of Sinorhizobium fredii NGR234 directly targets the plant’s symbiotic signaling pathway by associating with two Nod factor receptors (NFR1 and NFR5 of Lotus japonicus). NopT inhibits cell death triggered by co-expression of NFR1/NFR5 in Nicotiana benthamiana. Full-length NopT physically interacts with NFR1 and NFR5. NopT proteolytically cleaves NFR5 both in vitro and in vivo, but can be inactivated by NFR1 as a result of phosphorylation. NopT plays an essential role in mediating rhizobial infection in L. japonicus. Autocleaved NopT retains the ability to cleave NFR5 but no longer interacts with NFR1. Interestingly, genomes of certain Sinorhizobium species only harbor nopT genes encoding truncated proteins without the autocleavage site. These results reveal an intricate interplay between rhizobia and legumes, in which a rhizobial effector protease targets NFR5 to suppress symbiotic signaling. NFR1 appears to counteract this process by phosphorylating the effector. This discovery highlights the role of a bacterial effector in regulating a signaling pathway in plants and opens up the perspective of developing kinase-interacting proteases to fine-tune cellular signaling processes in general.

Introduction

The legume–rhizobium symbiosis stands as one of the most crucial and intricate mutualistic interactions in nature. Its significance goes beyond the supply of nitrogen fixed by rhizobia, which influences plant growth and ecosystem sustainability, as this symbiosis is a pivotal model for investigating plant-microbe interactions (Yang et al., 2022). The orchestration of this symbiosis involves the perception of rhizobial signal molecules such as lipo–chitooligosaccharidic Nod factors (NFs) and surface polysaccharides (Madsen et al., 2003; Tirichine et al., 2006; Kawaharada et al., 2015), along with the secretion of proteins, including type three protein secretion system (T3SS) effectors translocated into host plants (Sugawara et al., 2018; Zhang et al., 2021).

Earlier version: Plant receptor heterocomplexes consisting of two Lysin-motif receptor kinases (LYKs), play a central role in recognizing various microbial poly- and oligosaccharide molecules. In non–legumes such as Arabidopsis, AtLYK5 and AtCERK1 (Chitin elicitor receptor kinase) are required for perception of chitin, and thereby trigger plant immunity to combat fungal pathogens (Cao et al., 2014). In legumes, structurally similar Nod factor receptors (NFRs) such as NFR1 and NFR5 in Lotus japonicus are vital for recognition of rhizobial NFs and initiation of symbiotic signaling (Broghammer et al., 2012; Bozsoki et al., 2017). Accordingly, nfr1 and nfr5 knockout mutants are almost completely unable to form nodules (Madsen et al., 2003; Radutoiu et al., 2003). Over-expression of either NFR1 or NFR5 can activate NF signaling, resulting in formation of spontaneous nodules in the absence of rhizobia (Ried et al., 2014). The intricate action of NFRs is highlighted by the elictiation of a hypersensitive cell death response when NFR1/NFR5 are simultaneous over-expressed in Nicotiana benthamiana leaves (Madsen et al., 2011). Similarly, defense-like responses in Medicago truncatula nodules were observed upon over–expression of MtNFP (Nod factor perception, the ortholog of NFR5) (Moling et al., 2014), perhaps reflecting a potential role of MtNFP–interacting LYKs involved in plant–pathogen associations. It can therefore be expected that the protein abundance of NFRs must be tightly controlled to avoid resistance responses, thereby ensuring optimal root infection by rhizobia.

Revised version: Plant receptor heterocomplexes consisting of two Lysin-motif receptor kinases (LYKs), play a central role in recognizing various microbial poly- and oligosaccharide molecules. In non–legumes such as Arabidopsis, AtLYK5 and AtCERK1 (Chitin elicitor receptor kinase) are required for perception of chitin, and thereby trigger plant immunity to combat fungal pathogens (Cao et al., 2014). In legumes, structurally similar Nod factor receptors (NFRs) such as NFR1 and NFR5 in Lotus japonicus are vital for recognition of rhizobial NFs and initiation of symbiotic signaling (Broghammer et al., 2012; Bozsoki et al., 2017). Accordingly, nfr1 and nfr5 knockout mutants are almost completely unable to form nodules (Madsen et al., 2003; Radutoiu et al., 2003). Over-expression of either NFR1 or NFR5 can activate NF signaling, resulting in formation of spontaneous nodules in the absence of rhizobia (Ried et al., 2014). The intricate action of NFRs is highlighted by the elicitation of a hypersensitive cell death response when NFR1/NFR5 are simultaneous over-expressed in Nicotiana benthamiana leaves (Madsen et al., 2011). Similarly, defense-like responses in Medicago truncatula nodules were observed upon over–expression of MtNFP (Nod factor perception, the ortholog of NFR5) (Moling et al., 2014), perhaps reflecting a potential role of MtNFP–interacting LYKs involved in plant–pathogen associations. It can therefore be expected that the protein abundance of NFRs must be tightly controlled to avoid resistance responses, thereby ensuring optimal root infection by rhizobia.

Removed in the revised version: T3SS effectors of phytopathogenic Gram–negative bacteria play a critical role in suppressing pattern-triggered immunity (PTI) in host plants (Jones and Dangl, 2006). To avoid diseases, specific resistance (R) genes are employed to directly or indirectly recognize effectors translocated into host cells.

Earlier version: Recognition of such avirulence (Avr) effectors often results in effector-triggered immunity (ETI), a much stronger immune response which is usually associated with programmed cell death (Jones and Dangl, 2006). In contrast to the well–studied effector functions of phytopathogenic bacteria in suppressing immunity through various mechanisms, it is largely unknown how T3SS effectors from symbiotic rhizobia contribute to mutualistic interactions.

Revised version: T3SS effectors of phytopathogenic Gram–negative bacteria play a critical role in suppressing pattern-triggered immunity (PTI) in host plants (Jones and Dangl, 2006). To avoid diseases, specific resistance (R) genes are employed to directly or indirectly recognize effectors translocated into host cells. Recognition of such avirulence (Avr) effectors often results in effector-triggered immunity (ETI), a much stronger immune response which is usually associated with programmed cell death (Jones and Dangl, 2006). In contrast to the well–studied effector functions of phytopathogenic bacteria in suppressing immunity through various mechanisms, it is largely unknown how T3SS effectors from symbiotic rhizobia contribute to mutualistic interactions.

Earlier version: T3SS effectors of the YopT–type family, such as YopT of Yersinia pestis, AvrPphB of Pseudomonas syringae pv. phaseolicola and rhizobial NopT proteins, are evolutionary related cysteine proteases (Shao et al., 2002; Dai et al., 2008). YopT cleaves Rho family GTPases leading to their inactivation to disrupt the actin cytoskeleton of human host cells (Shao et al., 2003a). AvrPphB can proteolytically cleave Arabidopsis receptor–like cytoplasmic kinase family proteins including PBS1, PBS1-like proteins and BIK1, to suppress PTI (Shao et al., 2003b; Zhang et al., 2010). Cleavage of PBS1 kinase by AvrPphB can be regarded as a “decoy strategy” as it is associated with RPS5 (RESISTANCE TO PSEUDOMONAS SYRINGAE5)-mediated ETI (Shao et al., 2003b). Depending on the host legume, the NopT effector protease of Sinorhizobium fredii NGR234 can have an opposite function in regulating symbiosis (Dai et al., 2008; Kambara et al., 2009). The heterologous expression of NopT from NGR234 in S. fredii USDA 257 severely blocks nodulation in soybean Nanfeng 15 but not in other soybean cultivars (Khan et al., 2022), implying that NopT can act as an Avr effector. In contrast, NopT positively regulates nodule formation in the interaction between S. fredii HH103 and soybean Suinong 14 (Li et al., 2023). NopT from Bradyrhizobium sp. ORS3257 promotes rhizobial infection in nodules of Aeschynomene indica (Teulet et al., 2019). NopT proteases have been biochemically characterized. Some exhibit autocleavage activity and are subsequently acylated at the newly formed N-terminus (Dai et al., 2008; Dowen et al., 2009; Kambara et al., 2009; Fotiadis et al., 2012; Khan et al., 2022). Recently, soybean PBS1-1 and two proteins in Robinia pseudoacacia were identified as NopT targets (Luo et al., 2020; Khan et al. 2022; Li et al. 2023). However, at the molecular level, it remains largely unknown how NopT modulates rhizobial infection and whether NopTs from different rhizobial strains differ in their effector activities.

Revised version: T3SS effectors of the YopT–type family, such as YopT of Yersinia pestis, AvrPphB of Pseudomonas syringae pv. phaseolicola and rhizobial NopT proteins, are evolutionary related cysteine proteases (Shao et al., 2002; Dai et al., 2008). YopT cleaves Rho family GTPases leading to their inactivation to disrupt the actin cytoskeleton of human host cells (Shao et al., 2003a). AvrPphB can proteolytically cleave Arabidopsis receptor–like cytoplasmic kinase family proteins including PBS1, PBS1-like proteins and BIK1, to suppress PTI (Shao et al., 2003b; Zhang et al., 2010). Cleavage of PBS1 kinase by AvrPphB can be regarded as a “decoy strategy” as it is associated with RPS5 (RESISTANCE TO PSEUDOMONAS SYRINGAE 5)-mediated ETI (Shao et al., 2003b). Depending on the host legume, the NopT effector protease of Sinorhizobium fredii NGR234 can have an opposite function in regulating symbiosis (Dai et al., 2008; Kambara et al., 2009). The heterologous expression of NopT from NGR234 in S. fredii USDA257 severely blocks nodulation in soybean cultivar Nanfeng 15 but not in other soybean cultivars tested (Khan et al., 2022), implying that NopT might act as an Avr effector. In contrast, NopT positively regulates nodule formation in the interaction between S. fredii HH103 and soybean cultivar Suinong 14 (Li et al., 2023). NopT from Bradyrhizobium sp. ORS3257 promotes rhizobial infection in nodules of Aeschynomene indica (Teulet et al., 2019). NopT proteases have been biochemically characterized. Some exhibit autocleavage activity and are subsequently acylated at the newly formed N-terminus (Dai et al., 2008; Dowen et al., 2009; Kambara et al., 2009; Fotiadis et al., 2012; Khan et al., 2022). Recently, soybean PBS1-1 and two proteins in Robinia pseudoacacia were identified as NopT targets (Luo et al., 2020; Khan et al. 2022; Li et al. 2023). However, at the molecular level, it remains largely unknown how NopT modulates rhizobial infection and whether NopTs from different rhizobial strains differ in their effector activities.

Earlier version: A challenge in the investigation of rhizobial effectors lies in the availability of suitable assays and plants with which effector activities can be characterized. In this study, we took advantage of the observation that ectopic expression of NFR1 and NFR5 in N. benthamiana leaves induces programmed cell death (Madsen et al., 2011). We hypothesized that this response could be modulated by co-expressed rhizobial effectors. Among the 15 known or putative T3SS effectors of S. fredii NGR234, only NopT could specifically suppress NFR1/NFR5-induced cell death. This observation suggested that NopT is an effector associated with NFRs. Subsequent experiments demonstrated that NopT interacts with both NFR1 and NFR5 at the plasma membrane and proteolytically cleaves NFR5 at the juxtamembrane (JM) domain to suppress NF-mediated host responses. NFR1 phosphorylates the full-length NopT, thereby inactivating its protease activity. However, the autocleaved NopT form of S. fredii NGR234 retains the ability to cleave NFR5 but loses its ability to interact with NFR1. Intriguingly, various Sinorhizobium strains possess only truncated versions of NopT (related to autocleaved NopT of NGR234) which can evade the inactivation by NFR1. We present in this study a model of mutual regulation, in which the rhizobial effector NopT directly targets NFR5 to dampen symbiotic signaling, whereas NFR1-mediated phosphorylation of NopT counteracts NFR5 cleavage. Our work provides essential insights into the intricate interplay between legumes and rhizobia and shows an example of how a rhizobal effector fine-tunes NF signaling by directly targeting NFRs, and how an NFR inactivates the enzymatic activity of the effector through phosphorylation.

Revised version: A challenge in the investigation of rhizobial effectors lies in the availability of suitable assays and plants with which effector activities can be characterized. In this study, we took advantage of the observation that ectopic expression of NFR1 and NFR5 in N. benthamiana leaves induces programmed cell death (Madsen et al., 2011). We hypothesized that this response could be modulated by co-expressed rhizobial effectors. Among the 15 known or putative T3SS effectors of S. fredii NGR234, only NopT could suppress NFR1/NFR5-induced cell death. This observation suggested that NopT is an effector associated with NFRs. Subsequent experiments demonstrated that NopT interacts with both NFR1 and NFR5 at the plasma membrane and proteolytically cleaves NFR5 at the juxtamembrane (JM) domain to suppress NF-mediated host responses. NFR1 phosphorylates the full-length NopT, thereby inactivating its protease activity. However, the autocleaved NopT form of S. fredii NGR234 retains the ability to cleave NFR5 but loses its ability to interact with NFR1. Intriguingly, various Sinorhizobium strains possess only truncated versions of NopT (related to autocleaved NopT of NGR234) which might evade the inactivation by NFR1. We present in this study a model of mutual regulation, in which the rhizobial effector NopT directly targets NFR5 to dampen symbiotic signaling, whereas NFR1-mediated phosphorylation of NopT counteracts NFR5 cleavage. Our work provides essential insights into the intricate interplay between legumes and rhizobia and shows an example of how a rhizobal effector fine-tunes NF signaling by directly targeting NFRs, and how an NFR inactivates the enzymatic activity of the effector through phosphorylation.

Results

NopT expression in N. benthamiana suppresses the NFR1/NFR5-induced cell death response

Earlier version: To examine the role of effector proteins in regulating rhizobial infection, 15 effector genes (Kimbrel et al., 2013) were cloned from S. fredii NGR234, a strain with an exceptionally broad host-range. Initial experiments were performed to screen the effector proteins for their ability to suppress or enhance the NFR1/NFR5–triggered cell death response of N. benthamiana leaves. In this experiment, each effector was co-expressed with NFR1/NFR5 (Fig. S1A–S1C). Only NopT, a cysteine protease belonging to the YopT family (Shao et al., 2002; Dai et al., 2008), could suppress NFR1 and NFR5-induced cell death (Fig. 1A and S1D–S1F). Interestingly, NopT reduced protein abundance of the intact NFR5 when co–expressed with NFR1 (Fig. S1E), suggesting that NopT promotes proteolytic degradation of NFR5. Cys-93 of NopT is an amino acid residue known to be essential for the protease activity of this effector (Dai et al., 2008; Kambara et al., 2009). In contrast to the wild-type NopT protein, expression of the protease-dead NopTC93S form in N. benthamiana did not result in suppression of cell death induced by NFR1/NFR5 (Fig. 1B), indicating that the protease activity of NopT is essential for the observed cell death-suppressing activity of NopT. In contrast to NopT, NopM, an effector protein from S. fredii NGR234 reported to possess E3 ligase activity (Xin et al., 2012), induced cell death when expressed in N. benthamiana leaves under our test conditions (Fig. 1C).

Revised version: To examine the role of effector proteins in regulating rhizobial infection, 15 effector genes (Kimbrel et al., 2013) were cloned from S. fredii NGR234, a strain with an exceptionally broad host-range. Initial experiments were performed to screen the effector proteins for their ability to suppress or enhance the NFR1/NFR5–triggered cell death response of N. benthamiana leaves. In this experiment, each effector was co-expressed with NFR1/NFR5 (Fig. S1A–S1C). Only NopT, a cysteine protease belonging to the YopT family (Shao et al., 2002; Dai et al., 2008), could suppress NFR1 and NFR5-induced cell death (Fig. 1A and S1D–S1F). Interestingly, NopT reduced protein abundance of the intact NFR5 when co–expressed with NFR1 (Fig. S1E), suggesting that NopT might promote proteolytic degradation of NFR5. Cys-93 of NopT is an amino acid residue known to be essential for the protease activity of this effector (Dai et al., 2008; Kambara et al., 2009). In contrast to the wild-type NopT protein, expression of the protease-dead NopTC93S form in N. benthamiana did not result in suppression of cell death induced by NFR1/NFR5 (Fig. 1B), indicating that the protease activity of NopT is essential for the observed cell death-suppressing activity of NopT. In contrast to NopT, NopM, an effector protein from S. fredii NGR234 reported to possess E3 ligase activity (Xin et al., 2012), induced cell death when expressed in N. benthamiana leaves under our test conditions (Fig. 1C).

NopT specifically suppresses NFR1/NFR5-triggered cell death in N. benthamiana

Earlier version: Agrobacterium strains harboring plasmid DNA encoding NopT or the empty vector (EV) were infiltrated into N. benthamiana leaves. At 12 hpi, a second infiltration was performed with an Agrobacterium strain containing a plasmid with NFR1/NFR5 genes. Dashed red lines indicate leaf discs where different proteins were expressed. (A, B) NopT (A), but not NopTC93S (B), a protease-inactive version of NopT, suppressed NFR1NFR5-induced cell death. (C) Expression of NopM-induced cell death in N. benthamiana. (D) Expression of NopT could not suppress the cell death response triggered by expression of Avr3a/R3a, NopM, BAX, INF1 or AtCERK1. Numbers in (D) represent the number of leaf discs showing cell death and total leaf discs tested. The pictures shown in this figure are representative of at least three independent biological replicates.

Revised version: Agrobacterium strains harboring plasmid DNA encoding NopT or the empty vector (EV) were infiltrated into N. benthamiana leaves. At 12 hpi, a second infiltration was performed with an Agrobacterium strain containing a plasmid with NFR1/NFR5 genes. Dashed red lines indicate leaf discs where different proteins were expressed. (A, B) NopT (A), but not NopTC93S (B), a protease-inactive version of NopT, suppressed NFR1NFR5-induced cell death. (C) Expression of NopM-induced cell death in N. benthamiana. (D) Expression of NopT could not suppress the cell death response triggered by expression of Avr3a/R3a, NopM, BAX, INF1 or AtCERK1. Numbers in (D) represent the number of leaf discs showing cell death and total leaf discs tested. Cell death in leaf disc results in the formation of necrotic plaques, which restrains pathogens within deceased cells. These plaques commonly manifest as leaf dehydration, frequently accompanied by a translucent appearance. Brown and shriveled leaf discs serve as indicators of cell death. The pictures shown in this figure are representative of at least three independent biological replicates.

Earlier version: We also tested whether the cell death-suppressing effect of NopT was specific for NFR1/NFR5 expression or whether NopT could act generally to suppress cell death. In addition to NopM, we examined various proteins known to induce cell death when expressed in N. benthamiana leaves, namely Avr3a and R3a (Bos et al., 2006), apoptosis regulator BAX (Bcl-2-associated X protein) (Gan et al., 2009), INF1, an elicitor from Phytophthora infestans (Vleeshouwers et al., 2006) and Arabidopsis AtCERK1 (Cao et al., 2014) (Fig. 1C). In contrast to the NFR1/NFR5–induced cell death, expression of NopT was unable to suppress cell death when each of these proteins was expressed in N. benthamiana leaves (Fig. 1D). Hence, the effect of NopT appears to be specific for the NFRs. These results suggested that the effector protease NopT could act directly on NFR1 and/or NFR5.

Revised version: We also tested whether the cell death–suppressing effect of NopT was specific for NFR1/NFR5 expression or whether NopT could act generally to suppress cell death. In addition to NopM, we examined various proteins known to induce cell death when expressed in N. benthamiana leaves, namely Avr3a and R3a (Bos et al., 2006), apoptosis regulator BAX (Bcl-2-associated X protein) (Gan et al., 2009), INF1, an elicitor from Phytophthora infestans (Vleeshouwers et al., 2006) and Arabidopsis AtCERK1 (Cao et al., 2014) (Fig. 1C). In contrast to the NFR1/NFR5–induced cell death, expression of NopT was unable to suppress cell death when each of these proteins was expressed in N. benthamiana leaves (Fig. 1D). Hence, the effect of NopT appears to be specific for the NFRs. These results suggested that the effector protease NopT might act directly on NFR1 and/or NFR5.

NopT interacts with NFR1 and NFR5

Earlier version: Autocleaved NopT of S. fredii NGR234 contains acylation sites required for lipidation and subsequent plasma membrane localization (Dowen et al., 2009; Khan et al., 2022). To investigate whether NopT interacts with NFR1 and NFR5, we first used in vivo bimolecular fluorescence complementation (BiFC) analysis. NopT, NFR1 and NFR5 were C-terminally tagged with nYFP and cYFP. Fluorescence signals representing direct NopT-NFR1, NopT-NFR5 and NopT-NopT interactions were detected at the plasma membrane of N. benthamiana (Fig. 2A and 2B). Interactions between NopT and NFR1 and NFR5 were further verified using a Split-luciferase (Split–LUC) complementation assay (Fig. 2C). Co-expression of NopT-cLUC and NFR1-nLUC or NFR5-nLUC produced strong luminescence signals in transformed leaf discs compared with the negative control expressing AtFLS2-nLUC (Arabidopsis FLAGELLIN SENSING2) and NopT-cLUC (Fig. 2C).

Revised version: Autocleaved NopT of S. fredii NGR234 contains acylation sites required for lipidation and subsequent plasma membrane localization (Dowen et al., 2009; Khan et al., 2022). To investigate whether NopT interacts with NFR1 and NFR5, we first used in vivo bimolecular fluorescence complementation (BiFC) analysis. NopT, NFR1 and NFR5 were C-terminally tagged with nYFP and cYFP. Fluorescence signals representing direct NopT-NFR1, NopT-NFR5 and NopT-NopT interactions were detected at the plasma membrane of N. benthamiana (Fig. 2A and 2B). Interactions between NopT and NFR1 and NFR5 were further verified using a Split-luciferase (Split–LUC) complementation assay (Fig. 2C). Co-expression of NopT-CLuc and NFR1-NLuc or NFR5-NLuc produced strong luminescence signals in transformed leaf discs compared with the negative control expressing AtFLS2-NLuc (Arabidopsis FLAGELLIN SENSING 2) and NopT-CLuc (Fig. 2C).

NopT interacts with NFR1 and NFR5

Interactions between NopT and NFR1 or NFR5 were detected using BiFC (Split-YFP) (A, B), Split-LUC complementation (C) and co-IP (D) assays in N. benthamiana leaves. (A, B) For BIFC analysis, nYFP and cYFP tags were fused at the C-terminus of NopT, NFR1, NFR5, and the flagellin receptor FLS2 (negative control). YFP fluorescence signals represent protein-protein interactions. Scale bar=25 μm. (C) For the Split-LUC complementation assay, the NLuc and CLuc tags were fused at the C-terminus of NopT, NFR1, NFR5 and FLS2 (negative control). Luminescence signals represent protein-protein interactions. (D) For the co-IP assay, HA-tagged NFR1 or NFR5 and FLAG-tagged NopT were expressed in N. benthamiana cells followed by immunoprecipitation using an anti-FLAG antibody. Immunoblot analysis was performed using anti-HA and anti-FLAG antibodies (NopTC denotes the truncated version of NopT after autocleavage). The images and immunoblots shown in this figure are representative of three biological replicates

Earlier version: It is known that acylation of NopT can alter its subcellular localization (Dowen et al., 2009; Khan et al., 2022). Here, we performed Split-LUC assays to examine the interaction between NFR1/NFR5 and NopTG50A/C51A/C52A, a modified NopT form lacking acylation sites. Compared to the wild-type form (NopTWT), the interactions between NopTG50A/C51A/C52A and NFR1/NFR5 were significantly reduced in Split-LUC assays (Fig. S2A and S2B). However, the protease-dead NopT form (NopTC93S) showed interactions with NFR1 and NFR5 that were not different from NopTWT (Fig. S2C). These results were further confirmed by a co-immunoprecipitation (co-IP) experiment with N. benthamiana leaves co-expressing FLAG-tagged NopT with C-terminally HA-tagged NFR1 or NFR5. Both NFR1 and LjNFR5 could be co-precipitated by an anti-FLAG antibody when the sample contained NopT-FLAG (Fig. 2D). Overall, the identified physical interactions between NopT and NFR1/NFR5 indicated that NopT targets NFRs and that NFR5 and/or NFR1 may be proteolytically cleaved by NopT.

Revised version: It is known that acylation of NopT can alter its subcellular localization (Dowen et al., 2009; Khan et al., 2022). Here, we performed Split-LUC assays to examine the interaction between NFR1/NFR5 and NopTG50A/C51A/C52A, a modified NopT form lacking acylation sites. Compared to the wild-type form (NopTWT), the interactions between NopTG50A/C51A/C52A and NFR1/NFR5 were significantly reduced in Split-LUC assays (Fig. S3A and S3B). However, the protease-dead version of NopT (NopTC93S) showed interactions with NFR1 and NFR5, which were different from NopTWT (Fig. S3C). These results were further confirmed by a co-immunoprecipitation (co-IP) experiment in N. benthamiana leaves co-expressing FLAG-tagged NopT with C-terminally HA-tagged NFR1 or NFR5. Both NFR1 and NFR5 could be co-precipitated by an anti-FLAG antibody when the sample contained NopT-FLAG (Fig. 2D and S2). Overall, the identified physical interactions between NopT and NFR1/NFR5 indicated that NopT targets NFRs and that NFR5 and/or NFR1 may be proteolytically cleaved by NopT.

NopT proteolytically cleaves NFR5 at the JM domain

Earlier version: A previous study has shown that NopT is autocleaved at its N-terminus to form a processed protein that lacks the first 49 amino acid residues (Dai et al., 2008). To test whether NopT could proteolyze NFR1 and/or NFR5, NopT was co-expressed with NFR1-GFP and NFR-GFP fusion proteins in N. benthamiana leaves. The expressed proteins were separated by SDS-PAGE and subjected to immunoblotting. Proteolytic cleavage of NFR5-GFP but not NFR1-GFP was observed in the presence of co-expressed NopT, whereas protease-dead NopTC93S showed no effect (Fig. 3A and S3A). A similar in vivo cleavage assay with transgenic L. japonicus roots also showed that expression of NopT but not NopTC93S caused proteolytic cleavage of NFR5 (Fig. 3B). Based on the molecular weight changes, the cleavage site in NFR5 was predicted to occur within the cytoplasmic domain (CD) of NFR5. To further explore NFR5 cleavage by NopT, the CD of NFR5 (Strep-NFR5CD-HA) and NopT-FLAG or protease-dead NopTC93S-FLAG were co-expressed in Escherichia coli cells. Immunoblotting results with an anti-HA antibody showed that the active NopT protease cleaved NFR5CD resulting in a ∼5 kDa smaller protein. However, when the CD of NFR1 was co-expressed with NopT in the same experiment, no NFR1 cleavage was detected (Fig. 3C and 3D). Migration of the bands representing full-length NopT or NopTC93S was significantly retarded on the gel when the CD of NFR1 was co-expressed (Fig. 3C), suggesting that NopT was phosphorylated by NFR1 in E. coli. The NFR5CD cleavage product with a ∼5 kDa lower molecular weight in the cleavage assay suggested that the cleavage site is located within the juxtamembrane (JM) domain of NFR5. To test this hypothesis, we replaced the JM with the SUMO tag to creat SUMO-NFR5KC-HA (KC, kinase domain and C-terminal tail region, a modified NFR5CD without the JM) for proteolytic assay. Indeed, immunoblot analysis showed that co-expressed NopT was unable to cleave SUMO-NFR5KC-HA (Fig. S3B). The JM domain of NFR5 was then fused to SUMO and GFP to create a recombinant protein, SUMO-NFR5JM-GFP. In an in vitro cleavage assay with recombinant proteins from E. coli, GFP was immunodetected in the flow though sample when SUMO-NFR5JM-GFP was incubated with NopT but not with NopTC93S (Fig. 3E). The interaction between NopTC93S and SUMO-NFR5JM-GFP was confirmed by an in vitro pull-down assay (Fig. S4). Overall, these data strongly indicate that the cleavage site of NFR5 for NopT is located in its JM domain.

Revised version: A previous study has shown that NopT is autocleaved at its N-terminus to form a processed protein that lacks the first 49 amino acid residues (Dai et al., 2008). To test whether NopT could proteolyze NFR1 and/or NFR5, NopT was co-expressed with NFR1-GFP and NFR5-GFP fusion proteins in N. benthamiana leaves. The expressed proteins were separated by SDS-PAGE and subjected to immunoblotting. Proteolytic cleavage of NFR5-GFP but not NFR1-GFP was observed in the presence of co-expressed NopT, whereas protease-dead NopTC93S showed no effect (Fig. 3A and S4A). A similar in vivo cleavage assay with transgenic L. japonicus roots also showed that expression of NopT but not NopTC93S caused proteolytic cleavage of NFR5 (Fig. 3B). Based on the molecular weight changes, the cleavage site in NFR5 was predicted to occur within the cytoplasmic domain (CD) of NFR5. To further explore NFR5 cleavage by NopT, the CD of NFR5 (Strep-NFR5CD-HA) and NopT-FLAG or protease-dead NopTC93S-FLAG were co-expressed in the Escherichia coli cells. Immunoblotting results showed that the active NopT protease cleaved NFR5CD resulting in a ∼5 kDa smaller protein (Fig. 3C and 3D). However, when the CD of NFR1 was co-expressed with NopT in the same experiment, no NFR1 cleavage was detected (Fig. 3C). Migration of the bands representing full-length NopT or NopTC93S was significantly retarded on the gel when the CD of NFR1 was co-expressed (Fig. 3C), suggesting that NopT might be phosphorylated by NFR1 in E. coli cells.

NopT proteolyzes NFR5 at its JM domain

Earlier version: Proteins with indicated tags were expressed in N. benthamiana (A), L. japonicus (B) or E. coli (C-G) and detected by immunoblotting. (A) NopT but not NopTC93S cleaves NFR5-GFP protein expressed in N. benthamiana cells by releasing NFR5CD-GFP (NopTC denotes autocleaved NopT). (B) NopT but not NopTC93S cleaves NFR5-Myc expressed in hairy roots of L. japonicus by releasing NFR5CD-Myc (marked by an asterisk). (C) Analysis of the CDs of NFR1 and NFR5 co-expressed with NopT or NopTC93S in E. coli. Cleavage of NFR5CD was observed for NopT but not NopTC93S, while NFR1CD was not proteolyzed by NopT. In the presence of NFR1CD, a slower migrating band was observed, possibly representing phosphorylated NopT (NopTm). (D) A repeat experiment confirmed that NopT is able to cleave NFR5CD (the asterisk indicates the HA-tagged NFR5 cleavage product containing the kinase domain and a C-terminal tail region). (E) Not cleaves His-SUMO-NFR5JM-GFP (His-SUMO and GFP linked by the JM domain of NFR5) in vitro. F.T. indicates proteins in flow through samples after purification with Ni-beads. (F) NopT expressed in E. coli was unable to cleave co-expressed NFR5CD-5m-HA, a modified version of NFR5CD-HA in which five amino acids of the JM were substituted by other residues (S283Y, G294Q, Y303S, A310I, T311Y). (G) NopT expressed in E. coli was unable to cleave co-expressed M. truncatula NFPCD-HA and the NFR5JM-NFPKC-HA fusion protein, while NFPJM-NFR5KC was proteolyzed (KC stands for the kinase domain and a C-terminal tail region).

Revised version: Proteins with indicated tags were expressed in N. benthamiana (A), L. japonicus (B) or E. coli cells (C-G) and detected by immunoblotting. (A) NopT but not NopTC93S (a protease-dead version of NopT) cleaves NFR5-GFP protein expressed in N. benthamiana cells by releasing NFR5CD-GFP (NopTC denotes autocleaved NopT). The cleavage efficiency was marked under the lane. (B) NopT but not NopTC93S cleaves NFR5-Myc expressed in hairy roots of L. japonicus by releasing NFR5CD-Myc. The cleavage efficiency was marked under the lane. The asterisk indicates the HA-tagged NFR5 cleavage product containing the kinase domain and a C-terminal tail region (C) Analysis of the CDs of NFR1 and NFR5 co-expressed with NopT or NopTC93S in E. coli cells. Cleavage of NFR5CD was observed for NopT but not NopTC93S, while NFR1CD was not proteolyzed by NopT. In the presence of NFR1CD, a slower migrating band was observed, possibly representing phosphorylated NopT (NopTP). The cleavage efficiency was marked under the lane. (D) A repeat experiment confirmed that NopT is able to cleave NFR5CD (the asterisk indicates the HA-tagged NFR5 cleavage product containing the kinase domain and a C-terminal tail region). The cleavage efficiency was marked under the lane. (E) NopT cleaves His-SUMO-NFR5JM-GFP (His-SUMO and GFP linked by the JM domain of NFR5) in vitro. F.T. indicates proteins in flow through samples after purification with Ni-beads. (F) NopT expressed in E. coli cells was unable to cleave co-expressed NFR5CD-5m-HA, a modified version of NFR5CD-HA in which five amino acids of the JM were substituted by other residues (S283Y, G294Q, Y303S, A310I, T311Y). (G) NopT expressed in E. coli cells was unable to cleave co-expressed M. truncatula NFPCD-HA and the NFR5JM-NFPKC-HA fusion protein, while NFPJM-NFR5KC was proteolyzed (KC stands for the kinase domain and a C-terminal tail region, modified nod factor receptors without the JM).

Added in the revised version: The NFR5CD cleavage product with a ∼5 kDa lower molecular weight in the cleavage assay suggested that the cleavage site is located within the juxtamembrane (JM) domain of NFR5. To test this hypothesis, we replaced the JM with the SUMO tag to creat SUMO-NFR5KC-HA (KC, kinase domain and C-terminal tail region, a modified NFR5CD without the JM) for proteolytic assay. Indeed, immunoblot analysis showed that co-expressed NopT was unable to cleave SUMO-NFR5KC-HA (Fig. S4B). The JM domain of NFR5 was then fused to SUMO and GFP to create a recombinant protein, SUMO-NFR5JM-GFP. In an in vitro cleavage assay with recombinant proteins from E. coli cells, GFP was immunodetected in the flow though sample when SUMO-NFR5JM-GFP was incubated with NopT but not with NopTC93S (Fig. 3E). The interaction between NopTC93S and SUMO-NFR5JM-GFP was confirmed by an in vitro pull-down assay (Fig. S5). Overall, these data strongly indicate that the cleavage site of NFR5 for NopT is located in its JM domain.

The cleavage of NFR5 by NopT is dependent on multiple residues of the JM domain

Earlier version: The “DKLLSGV” motif in the JM domain of NFR5 (residues Asp-288 to Val-294; Fig. S5) shows highest similarity with the autocleavage region of NopT (DKMGCCA). However, a protein variant of NFR5CD in which the “DKLLSGV” motif was replaced by 7 alanine residues could still be cleaved by NopT (Fig. S6A). Similar to NopT, the AvrPphB effector of P. syringa has an autocleavage site (Shao et al., 2003b). We therefore also examined a version of NFR5CD in which the “DKLLSGV” motif was replaced by that of AvrPphB. In the cleavage assay with E. coli cells, AvrPphB did not cleave NFR5CD. However, the modified NFR5CD form with the 7 residues from AvrPphB was cleaved by AvrPphB and the cleavage product had a lower molecular weight than that formed by NopT (Fig. S6A). These findings indicated that the cleavage site for NopT is located upstream of the “DKLLSGV” motif in NFR5.

Revised version: The “DKLLSGV” motif in the JM domain of NFR5 (residues Asp-288 to Val-294; Fig. S6) shows highest similarity with the autocleavage region of NopT (DKMGCCA). However, a protein variant of NFR5CD in which the “DKLLSGV” motif was replaced by 7 alanine residues could still be cleaved by NopT (Fig. S7A). Similar to NopT, the AvrPphB effector of P. syringa has an autocleavage site (Shao et al., 2003b). We therefore also examined a version of NFR5CD in which the “DKLLSGV” motif was replaced by that of AvrPphB. In the cleavage assay with E. coli cells, AvrPphB did not cleave NFR5CD. However, the modified NFR5CD form with the 7 residues from AvrPphB was cleaved by AvrPphB and the cleavage product had a lower molecular weight than that formed by NopT (Fig. S7A). These findings indicated that the cleavage site for NopT is located upstream of the “DKLLSGV” motif in NFR5.

Earlier version: As mentioned above, NopT undergoes autocleavage but it is not clear whether this is due to intermolecular or intramolecular proteolysis. In addition to Cys-93, His-205 and Aps-220 of NopT are conserved catalytic residues (Dai et al., 2008; Kambara et al. 2009). Therefore, NopT forms lacking these catalytic residues were co-expressed with NopT in E. coli. The immunoblot analysis showed that NopT could not proteolyze NopTC93S, NopTH205A or NopTD220A (Fig. S6B), suggesting that NopT autocleavage is due to intramolecular proteolysis. Hence, the residues at the autocleavage site of NopT may not help to predict the cleavage site of NopT substrates.

Revised version: As mentioned above, NopT undergoes autocleavage but it is not clear whether this is due to intermolecular or intramolecular proteolysis. In addition to Cys-93, His-205 and Aps-220 of NopT are conserved catalytic residues (Dai et al., 2008; Kambara et al. 2009). Therefore, NopT forms lacking these catalytic residues were co-expressed with NopT in E. coli cells. The immunoblot analysis showed that NopT could not proteolyze NopTC93S, NopTH205A or NopTD220A (Fig. S7B), suggesting that NopT autocleavage is due to intramolecular proteolysis. Hence, the residues at the autocleavage site of NopT may not help to predict the cleavage site of NopT substrates.

Earlier version: In an attempt to define the NopT cleavage site in NFR5, we created a series of modified NFR5 forms covering the JM domain of NFR5 (Val-269 to Cys-320) and the first seven amino acid residues of the kinase domain (Lys-321 to Tyr-327). Initially, 17 variants of NFR5CD were created in which three adjacent amino acids in the JM domain were replaced with three alanine residues. When these proteins were co-expressed with NopT in E. coli, all 17 NFR5CD variants could be cleaved by NopT but not NopTC93S (Fig.S6C-S6F). In a similar experiment, we created seven NFR5CD forms which contained a deletion in the JM domain of 10 residues each. Surprisingly, cleavage products were still observed for all seven examined proteins (Fig. S6G). These results suggest that the JM domain may have multiple sites that can be cleaved by NopT, which is different from the specific cleavage site identified for AvrPphB. In another experiment, recombinant His-SUMO-NFR5CD-GST protein was co-expressed with NopT in E. coli and its cleavage product was subjected to N-terminal protein sequencing using Liquid Chromatography Mass Sprectrometry (LC-MS) analysis. Based on alignment of identified peptides, the NopT cleavage site in NFR5 was mapped to four basic amino acids in the JM domain (RRKK; amino acid residues 271-275) (Fig. S7). However, this result was not consistent with the observation that NopT expressed in E. coli could still proteolyze a co-expressed NFR5CD form in which the residues from Tyr-268 to Leu-277 were deleted (Fig. S6G). Overall, the experiments performed suggest that NopT preferentially proteolyzes NFR5 at the RRKK motif, but other sites in the JM domain can also be cleaved by NopT. Since NFR5 belongs to a subfamily of LYK family proteins that lacks kinase activity (Yang et al., 2022), we wondered whether NopT could proteolyze other LYK proteins of this subfamily, which might be helpful to characterize the biochemical function of NopT. We therefore investigated whether the CDs from Arabidopsis AtLYK5, L. japonicus LjLYS11 and M. truncatula MtNFP expressed in E. coli can be proteolyzed by NopT. As shown in Fig. S8, NopT, but not NopTC93S, was able to cleave the CDs of AtLYK5 and LjLYS11. Replacement of the JM domain in NFR5CD with the JM domain from either AtLYK5 or LjLYS11 resulted in cleavage by NopT (Fig. S8). These data show that NopT cleaves AtLYK5 and LjLYS11 at their JM domains. Based on sequence analysis of the JM region of NFR5 and its homologous proteins, five conserved residues (Ser-283, Gly-294, Tyr-304, Ala-310, Thr-311) were identified and mutated as NFR5CD-5m for the proteolytic assay (Fig. S5B). As shown in Fig. 3F, NFR5CD-5m, a NFR5CD variant with point mutations of S283Y, G294Q, Y304S, A310I, and T311Y, could not be cleaved by NopT. Interestingly, NopT could not proteolyze MtNFPCD (Fig. 3G). We then swapped the JM domains in NFR5CD and MtNFPCD, generating NFPJM-NFR5KC and NFR5JM-NFPKC (KC stands for the kinase domain and the C-terminal tail region). As shown in Fig. 3G, NopT could proteolyze NFPJM-NFR5KC but not NFR5JM-NFPKC in E. coli. In an in vitro cleavage assay, the cleavage product of SUMO-NFPJM-GFP proteolyzed by NopT was detected in the flow though sample (Fig. S9). These data suggest that the cleavage of NFR5 by NopT depends on the conformational structure of NFR5KC and/or other specific amino acids in NFR5KC. We then investigated whether two additional recombinant proteins with large swapping regions in NFR5CD and NFPCD (NFR5268-445-NFP458-595 and NFP270-457-NFR5456-595) could be cleaved in the E. coli proteolysis test (Fig. S10). Notably, both proteins were cleaved by NopT, suggesting that a specific conformation of the NopT substrate is required for proteolysis. Taken together, we concluded from our experiments that NopT can proteolyze NFR5 at the JM domain and that mutiple nonadjacent residues in NFR5JM in combination with a specific KC conformation are required for proteolysis.

Revised version: In an attempt to define the NopT cleavage site in NFR5, we created a series of modified NFR5 forms covering the JM domain of NFR5 (Val-269 to Cys-320) and the first seven amino acid residues of the kinase domain (Lys-321 to Tyr-327). Initially, 17 variants of NFR5CD were created in which three adjacent amino acids in the JM domain were replaced with three alanine residues. When these proteins were co-expressed with NopT in E. coli cells, all 17 NFR5CD variants could be cleaved by NopT but not NopTC93S (Fig.S7C-S7F). In a similar experiment, we created seven NFR5CD forms which contained a deletion in the JM domain of 10 residues each. Surprisingly, cleavage products were still observed for all seven examined proteins (Fig. S7G). These results suggest that the JM domain may have multiple sites that can be cleaved by NopT, which is different from the specific cleavage site identified for AvrPphB. In another experiment, recombinant His-SUMO-NFR5CD-GST protein was co-expressed with NopT in E. coli cells and its cleavage product was subjected to N-terminal protein sequencing using Liquid Chromatography Mass Sprectrometry (LC-MS) analysis. Based on alignment of identified peptides, the NopT cleavage site in NFR5 was mapped to four basic amino acids in the JM domain (RRKK; amino acid residues 271-275) (Fig. S8). However, this result was not consistent with the observation that NopT expressed in E. coli cells could still proteolyze a co-expressed NFR5CD form in which the residues from Tyr-268 to Leu-277 were deleted (Fig. S7G). Overall, the experiments performed suggest that NopT preferentially proteolyzes NFR5 at the RRKK motif, but other sites in the JM domain can also be cleaved by NopT. Since NFR5 belongs to a subfamily of LYK family proteins that lacks kinase activity (Yang et al., 2022), we wondered whether NopT could proteolyze other LYK proteins of this subfamily, which might be helpful to characterize the biochemical function of NopT. We therefore investigated whether the CDs from Arabidopsis AtLYK5, L. japonicus LjLYS11 and M. truncatula MtNFP expressed in E. coli cells can be proteolyzed by NopT. As shown in Fig. S9, NopT, but not NopTC93S, was able to cleave the CDs of AtLYK5 and LjLYS11. Replacement of the JM domain in NFR5CD with the JM domain from either AtLYK5 or LjLYS11 resulted in cleavage by NopT (Fig. S9). These data show that NopT cleaves AtLYK5 and LjLYS11 at their JM domains. Based on sequence analysis of the JM region of NFR5 and its homologous proteins, five conserved residues (Ser-283, Gly-294, Tyr-304, Ala-310, Thr-311) were identified and mutated as NFR5CD-5m for the proteolytic assay (Fig. S6B). As shown in Fig. 3F, NFR5CD-5m, a NFR5CD variant with point mutations of S283Y, G294Q, Y304S, A310I, and T311Y, could not be cleaved by NopT. Interestingly, NopT could not proteolyze MtNFPCD (Fig. 3G). We then swapped the JM domains in NFR5CD and MtNFPCD, generating NFPJM-NFR5KC and NFR5JM-NFPKC (KC stands for the kinase domain and the C-terminal tail region). As shown in Fig. 3G, NopT could proteolyze NFPJM-NFR5KC but not NFR5JM-NFPKC in E. coli cells. In an in vitro cleavage assay, the cleavage product of SUMO-NFPJM-GFP proteolyzed by NopT was detected in the flow though sample (Fig. S10). These data suggest that the cleavage of NFR5 by NopT depends on the conformational structure of NFR5KC and/or other specific amino acids in NFR5KC. We then investigated whether two additional recombinant proteins with large swapping regions in NFR5CD and NFPCD (NFR5268-445-NFP458-595 and NFP270-457-NFR5456-595) could be cleaved in the E. coli proteolysis test (Fig. S11). Notably, both proteins were cleaved by NopT, suggesting that a specific conformation of the NopT substrate is required for proteolysis. Taken together, we concluded from our experiments that NopT can proteolyze NFR5 at the JM domain and that mutiple nonadjacent residues in NFR5JM in combination with a specific KC conformation are required for proteolysis.

NopT phosphorylated by NFR1 is proteolytically inactive

The perception of rhizobial NFs by the heterocomplex of NFR1 and NFR5 involves trans-phosphorylation events from the NFR1 kinase to the pseudo-kinase NFR5 (Madsen et al., 2011). Since NopT is associated with both NFR1 and NFR5, we hypothesized that NopT is a phosphorylation target of NFR1, which may interfere with transphosphorylation events between the two receptors. NopTC93S was therefore expressed in L. japonicus wild-type and NFR1 knockout mutant plants (nfr1-1) and analyzed on Zn2+-Phos-tag gels. On such gels, a Zn2+-Phos-tag bound phosphorylated protein migrates slower than its unbound nonphosphorylated form. Slower migrating bands corresponding to phosphorylated NopTC93S were observed, particularly when wild-type plants were inoculated with rhizobia (Fig. 4A and 4B). These findings indicated that NopT phosphorylation in planta is largely dependent on NFR1 and that the protein was probably phosphorylated at multiple sites.

Phosphorylation of NopT by NFR1 suppresses its proteolytic activity

Earlier version: (A) In vivo phosphorylation assay with proteins expressed in L. japonicus roots (wild-type and nfr1-1 mutant plants) using Zn2+-Phos-tag SDS-PAGE. Phosphorylation of NopTC93S was induced by inoculation with rhizobia (Mesorhizobium loti MAFF303099) and was largely dependent on NFR1. (B) The relative protein amount of each lane, as shown in (A), was quantified with ImageJ software (three biological replicates). The value of the control band in each gel was set to 1 for comparison. Values are means ± SEM. (C) NFR1CD but not NFR1CD-K351E phosphorylates NopT in E. coli. The phosphorylated full-length form of NopT could be dephosphorylated by CIAP (gel band shift). Abbreviations: NFR1CD-P, autophosphorylated NFR1; NopTP, phosphorylated NopT; NopT, non-phosphorylated NopT; NopTC, autocleaved NopT. (D) The phoshorylation sites of NopT identified by LC-MS were either substituted to alanine (A) or aspartate (D). The indicated NopT variants were subsequently tested for autocleavage and NFR5CD proteolysis in E. coli. Wild-type NopT (WT) and protease-dead NopTC93S were included into the analysis.

Revised version: (A) In vivo phosphorylation assay with proteins expressed in L. japonicus roots (wild-type and nfr1-1 mutant plants) using Zn2+-Phos-tag SDS-PAGE. Phosphorylation of NopTC93S was induced by inoculation with rhizobia (Mesorhizobium loti MAFF303099) and was largely dependent on NFR1. (B) The relative protein amount of each lane, as shown in (A), was quantified with ImageJ software (three biological replicates). The value of the control band in each gel was set to 1 for comparison. Values are means ± SEM. (C) NFR1CD but not NFR1CD-K351E phosphorylates NopT in E. coli cells. The phosphorylated full-length form of NopT could be dephosphorylated by CIAP (gel band shift). Abbreviations: NFR1CD-P, autophosphorylated NFR1; NopTP, phosphorylated NopT; NopT, non-phosphorylated NopT; NopTC, autocleaved NopT. (D) The phosphorylation sites of NopT identified by LC-MS were either substituted to alanine (A) or aspartate (D). The indicated NopT variants were subsequently tested for autocleavage and NFR5CD proteolysis in E. coli cells. Wild-type NopT (WT) and protease-dead NopTC93S were included into the analysis.

Earlier version: Next, we performed assays to investigate whether expressed NFR1CD directly phosphorylates co-expressed NopT. NFR5CD and a kinase-inactive NFR1CD form (NFR1CD-K351E; Lys-351 is a conserved residue required for ATP binding) were included into these experiments. When NFR1CD was expressed in E. coli, extracted NopT (Fig. 3C and 4C) and NFR5CD (Fig. S11) exhibited retarded migration on SDS-PAGE gels, indicating that NFR1 was able to phosphorylate both proteins. A band shift was observed when the phosphorylated full-length NopT was in vitro dephosphorylated by calf intestinal alkaline phosphatase (CIAP). However, the CIAP treatment caused no obvious band shift of the autocleaved NopT form (Fig. 4C). The phosphorylation sites of NopT were then identified using an in vitro phosphorylation assay followed by LC-MS analysis. Several serine and threonine residues in NopT were identified to be phosphorylated by NFR1CD (Table S1). Taken together, these results showed that full-length NopT is a phosphorylation target of the NFR1 kinase. Finally, we wondered whether NopT phosphorylation by NFR1 influences its proteolytic activity. Plasmids encoding NopT variants were constructed in which the phosphorylated residues were substituted to either alanine (to block phosphorylation) or aspartate (to mimic the phosphorylation status). The different NopT forms were then expressed in E. coli to analyze NopT autocleavage and proteolysis of co-expressed NFR5CD. Aspartate substitutions at three phosphorylation sites of NopT (Thr-97, Ser-119 and Thr-257) resulted in the loss of the autoproteolytic activity, as well as the ability to cleave+ NFR5CD (Fig. 4D). In contrast, corresponding NopT proteins with alanine substitutions retained autocleavage activity and the ability to proteolyze NFR5CD (Fig. 4D). These data indicate that phosphorylation of full-length NopT by NFR1 inhibits its proteolytic activity, thereby keeping NopT unprocessed and NFR5 uncleaved.

Revised version: Next, we performed assays to investigate whether expressed NFR1CD directly phosphorylates co-expressed NopT. NFR5CD and a kinase-inactive NFR1CD form (NFR1CD-K351E; Lys-351 is a conserved residue required for ATP binding) were included into these experiments. When NFR1CD was expressed in E. coli cells, extracted NopT (Fig. 3C and 4C) and NFR5CD (Fig. S12) exhibited retarded migration on SDS-PAGE gels, suggesting that NFR1 was able to phosphorylate both proteins. A band shift was observed when the phosphorylated full-length NopT was in vitro dephosphorylated by calf intestinal alkaline phosphatase (CIAP). However, the CIAP treatment caused no obvious band shift of the autocleaved NopT form (Fig. 4C). The phosphorylation sites of NopT were then identified using an in vitro phosphorylation assay followed by LC-MS analysis. Several serine and threonine residues in NopT were identified to be phosphorylated by NFR1CD (Table S1). Taken together, these results showed that full-length NopT is a phosphorylation target of the NFR1 kinase.

Added in the revised version: Finally, we wondered whether NopT phosphorylation by NFR1 influences its proteolytic activity. Plasmids encoding NopT variants were constructed in which the phosphorylated residues were substituted to either alanine (to block phosphorylation) or aspartate (to mimic the phosphorylation status). The different NopT forms were then expressed in E. coli cells to analyze NopT autocleavage and proteolysis of co-expressed NFR5CD. Aspartate substitutions at three phosphorylation sites of NopT (Thr-97, Ser-119 and Thr-257) resulted in the loss of the autoproteolytic activity, as well as the ability to cleave NFR5CD (Fig. 4D). In contrast, corresponding NopT proteins with alanine substitutions retained autocleavage activity and the ability to proteolyze NFR5CD (Fig. 4D). These data indicate that phosphorylation of full-length NopT by NFR1 inhibits its proteolytic activity, thereby keeping NopT unprocessed and NFR5 uncleaved.

NopT dampens rhizobial infection

Earlier version: As NopT of S. fredii NGR234 could directly target and cleave NFR5, we expected that NopT may be an important player regulating rhizobial infection in legumes. NF signaling is known to function during early stages of rhizobial infection (Geurts et al., 2005; Wang et al., 2012; Cai et al., 2018). We focused our experiments on early stages of rhizobial infection. Inoculation of L. japonicus plants with a GFP- or GUS-labelled nopT mutant of NGR234 (NGR234 ΔnopT) resulted in a massive infection of root hairs, whereas infection by the NGR234 wild-type strain was less frequent (Fig. 5A and 5C). We also performed experiments with L. japonicus plants containing pNIN:GUS which exhibit β-glucuronidase (GUS) activity when the NIN (NODULE INCEPTION) promoter activity is induced by NF signaling (Schauser et al., 1999). GUS staining of roots showed that the NIN promoter activity was stronger in response to NGR234ΔnopT inoculation relative to inoculatioin by wild-type NGR234 (Fig. 5B and 5D). Likewise, inoculation with NGR234ΔnopT resulted in an increased number of nodule primordia (Fig. 5G). In contrast, over-expression of Not under the control of the T7 promoter in the NGR234 parent strain reduced the number of nodule primordia when compared to inoculation with the wild-type strain (Fig. 5G). These findings indicate that NopT negatively influences symbiosis at early symbiotic stages. However, we could not find any uninfected (empty) nodules when roots were inoculated with these strains.

Revised version: As NopT of S. fredii NGR234 could directly target and cleave NFR5, we expected that NopT may be an important player regulating rhizobial infection in legumes. NF signaling is known to function during early stages of rhizobial infection (Geurts et al., 2005; Wang et al., 2012; Cai et al., 2018). We focused our experiments on early stages of rhizobial infection. Inoculation of L. japonicus plants with a GFP-or GUS-labelled nopT mutant of NGR234 (NGR234 ΔnopT) resulted in a massive infection of root hairs, whereas infection by the NGR234 wild-type strain was less frequent (Fig. 5A and 5C). We also performed experiments with L. japonicus plants containing pNIN:GUS which exhibit β-glucuronidase (GUS) activity when the NIN (NODULE INCEPTION) promoter activity is induced by NF signaling (Schauser et al., 1999). GUS staining of roots showed that the NIN promoter activity was stronger in response to NGR234ΔnopT inoculation relative to inoculation by wild-type NGR234 (Fig. 5B and 5D). Likewise, inoculation with NGR234ΔnopT resulted in an increased number of nodule primordia (Fig. 5G). In contrast, over-expression of NopT under the control of the T7 promoter in the NGR234 parent strain reduced the number of nodule primordia when compared to inoculation with the wild-type strain (Fig. 5G). These findings indicate that NopT negatively influences rhizobial infection and nodule initiation. Interestingly, the internal structures of nodules stained by toluidine blue dye exhibited remarkable similar when inoculated with both NGR234 and NGR234ΔnopT (Fig. S13).

NopT regulates rhizobial infection in L. japonicus

Earlier version: (A) Analysis of rhizobial infection in L. japonicus roots inoculated with GUS-labelled S. fredii NGR234 (wild-type; WT) or a nopT knockout mutant (NGR234ΔnopT; abbreviated as ΔnopT in other panels). Scale bar=100 µm. (B) GUS staining pictures showing roots of L. japonicus plants expressing GUS expression under control of the NIN promoter (pNIN:GUS). Plants were inoculated with NGR234 or NGR234ΔnopT and analyzed at 7 dpi. Scale bar=1 mm. (C) Infection data (ITs, infection threads) for roots shown in (A) at 7 dpi (n=10, Student’s t-test; * indicates P < 0.01). (D) Quantification of GUS staining sites for roots shown in (B) (n=10, Student’s t-test, P < 0.01). (E) Nodule primordia formation in L. japonicus wild-type roots inoculated with NGR234 (WT), NGR234ΔnopT or NGR234 over-expressing nopT (pT7:NopT). Roots were analyzed at 14 dpi (n=16, Student’s t-test: P < 0.01). (F) Infection data for wild-type roots inoculated with NGR234 (WT), NGR234ΔnopT and NGR234ΔnopT expressing indicated NopT variants at 7 dpi (n=5, Student’s t-test: P < 0.01). (G) Analysis of rhizobial infection in hairy roots of L. japonicus (wild-type) expressing GFP (EV, empty vector control) or NopT. Plants were inoculated with DsRed-labeled M. loti MAFF303099 and analyzed at 5 dpi (n=8, Student’s t–test: P < 0.01). (H) Expression of NFR5 and NFR5-NFPKC in hairy roots of nfr5-3 mutant plants. Plants were inoculated with GFP-labelled NGR234 (WT) or NGR234ΔnopT and analyzed at 8 dpi. Roots expressing NFR5-NFPKC showed high numbers of infection foci for both strains whereas significant differences were observed for NFR5 expressing roots (n=10, Student’s t-test: P < 0.01).

Revised version: (A) Analysis of rhizobial infection in L. japonicus roots inoculated with GUS-labelled S. fredii NGR234 (wild-type; WT) or a nopT knockout mutant (NGR234ΔnopT; abbreviated as ΔnopT in other panels). Scale bar=100 µm. Infection represents both infection focus and infection thread. The infection threads in the root hairs were shown in the left image in the upper pannel, while the infection foci were shown in the left image in the middle pannel. (B) GUS staining pictures showing roots of L. japonicus plants expressing GUS expression under control of the NIN promoter (pNIN:GUS). Plants were inoculated with NGR234 or NGR234ΔnopT and analyzed at 7 dpi. Scale bar=1 mm. (C) Infection data (ITs, infection threads) for roots shown in (A) at 7 dpi (n=10, Student’s t-test; * indicates P < 0.01). (D) Quantification of GUS staining sites for roots shown in (B) (n=10, Student’s t-test, P < 0.01). (E) Nodule primordia formation in L. japonicus wild-type roots inoculated with NGR234 (WT), NGR234ΔnopT or NGR234 over-expressing nopT (pT7:NopT). Roots were analyzed at 14 dpi (n=16, Student’s t-test: P < 0.01). (F) Infection data for wild-type roots inoculated with NGR234 (WT), NGR234ΔnopT and NGR234ΔnopT expressing indicated NopT variants at 7 dpi (n=5, Student’s t-test: P < 0.01). (G) Analysis of rhizobial infection in hairy roots of L. japonicus (wild-type) expressing GFP (EV, empty vector control) or NopT. Plants were inoculated with DsRed-labeled M. loti MAFF303099 and analyzed at 5 dpi (n=8, Student’s t–test: P < 0.01). (H) Expression of NFR5 and NFR5-NFPKC in hairy roots of nfr5-3 mutant plants. Plants were inoculated with GFP-labelled NGR234 (WT) or NGR234ΔnopT and analyzed at 8 dpi. Roots expressing NFR5-NFPKC showed high numbers of infection foci for both strains whereas significant differences were observed for NFR5 expressing roots (n=10, Student’s t-test: P < 0.01).

Earlier version: To confirm that the massive infection of roots by NGR234ΔnopT is due to the loss of NopT, we complemented the NGR234ΔnopT by expressing wild-type NopT. The L. japonicus plants inoculated with this complemented strain had fewer infection foci, indicating restoration of the wild-type phenotype (Fig. 5E). We also tested whether the proteolytic activity of NopT affected rhizobial infection. The protease-dead version of NopTC93S and three NopT forms with phospho-mimimetic residues (i.e., NopTT97D, NopTS119D, and NopTT257D lacking proteolytic activity) were individually expressed in NGR234ΔnopT. Compared with the high number of infection foci induced by NGR234ΔnopT in L. japonicus plants, inoculation with these strains resulted in significantly fewer infection foci, but more than with the wild-type strain (Fig. 5E).

Revised version: To confirm that the massive infection of roots by NGR234ΔnopT is due to the loss of NopT, we complemented the NGR234ΔnopT by expressing wild-type NopT. The L. japonicus plants inoculated with this strain had fewer infection foci, indicating restoration of the wild-type phenotype (Fig. 5E). We also tested whether the proteolytic activity of NopT affected rhizobial infection. The protease-dead version of NopTC93S and three NopT forms with phospho-mimetic residues (i.e., NopTT97D, NopTS119D, and NopTT257D lacking proteolytic activity) were individually expressed in NGR234ΔnopT. Compared with the high number of infection foci induced by NGR234ΔnopT in L. japonicus plants, inoculation with these strains resulted in significantly fewer infection foci, but more than with the wild-type strain (Fig. 5E).

We also expressed nopT under the control of a ubiquitin promoter in hairy roots of L. japonicus and inoculated the transgenic roots with DsRed-labeled M. loti MAFF303099, which does not possess a nopT gene in its genome. As shown in Fig. 5F, reduced numbers of infection foci and infection threads were observed in the NopT expressing roots, indicating that NopT negatively affects the symbiosis between L. japonicus and M. loti strain.

Earlier version: To investigate whether the cleavage of NFR5 by NopT reduces rhizobial infection of L. japonicus roots, NFR5 and NFR55m (uncleavable variant with 5 amino acid substitutions in the JM domain; Fig. 3F) under the control of the native NFR5 promoter, were expressed in hairy roots of NFR5 knockout mutant plants (nfr5-3). NFR5 expression resulted in numerous infections and a significant increase of infection foci was observed upon inoculation with the NGR234ΔnopT mutant (Fig. 5H). However, NFR55m expression showed no effects on rhizobial infection (Fig. S12), suggesting that the amino acid substitutions in NFR55m caused a conformational change that rendered the protein inactive in symbiotic signaling. As the NFR5JM-NFPKC fusion protein was not cleaved by NopT (Fig. 3G), we also constructed hairy roots of nfr5-3 mutant plants in which NFR5-NFPKC (NFR5 variant with its KC domain replaced by the KC of MtNFP) was expressed under the control of the NFR5 promoter. Remarkably, inoculation of these plants showed no difference between NGR234ΔnopT and the NGR234 wild-type strain in terms of rhizobial infection foci (Fig. 5H), suggesting that the cleavage of NFR5 by NopT reduced the degree of rhizobial infection. Taken together, these data indicate that the protease activity of NopT, its phosphorylation status and the proteolysis of the NFR5 target are important for the regulation of rhizobial infection.

Revised version: To investigate whether the cleavage of NFR5 by NopT reduces rhizobial infection of L. japonicus roots, NFR5 and NFR55m(uncleavable variant with 5 amino acid substitutions in the JM domain; Fig. 3F) under the control of the native NFR5 promoter, were expressed in hairy roots of NFR5 knockout mutant plants (nfr5-3). NFR5 expression resulted in numerous infections and a significant increase of infection foci was observed upon inoculation with the NGR234ΔnopT mutant (Fig. 5H). However, NFR55m expression showed no effects on rhizobial infection (Fig. S14), suggesting that the amino acid substitutions in NFR55m caused a conformational change that rendered the protein inactive in symbiotic signaling. As the NFR5JM-NFPKC fusion protein was not cleaved by NopT (Fig. 3G), we also constructed hairy roots of nfr5-3 mutant plants in which NFR5-NFPKC (NFR5 variant with its KC domain replaced by the KC of MtNFP) was expressed under the control of the NFR5 promoter. Remarkably, inoculation of these plants showed no difference between NGR234ΔnopT and the NGR234 wild-type strain in terms of rhizobial infection foci (Fig. 5H), suggesting that the cleavage of NFR5 by NopT reduced the degree of rhizobial infection. Taken together, these data indicate that the protease activity of NopT, its phosphorylation status and the proteolysis of the NFR5 target are important for the regulation of rhizobial infection.

NopT from other S. fredii strains also cleave NFR5

Earlier version: Since NopT of S. fredii NGR234 (NopTNGR234) cleaves NFR5, we wondered whether homologs from other rhizobial species also possess such a proteolytic activity. Phylogenetic analysis showed that NopT proteins of Sinorhizobium and Bradyrhizobium are located in different clades (Fig. S13). B. diazoefficiens USDA110, a typical Bradyrhizobium strain used to study nodulation of soybeans, produces two NopT proteins (Fotiadis et al., 2012). In contrast to NopTNGR234, however, neither NopT1USDA110 nor NopT2USDA110 were able to cleave co-expressed NFR5 in E. coli (Fig. 6A). This finding suggests that NopT homologs from Bradyrhizbium species have lost the ability to cleave NFR5 and probably act on other, unknown host target proteins. The genome of various Sinorhizobium species (e.g., S. fredii USDA257 and S. fredii HH103) possess genes that encode truncated NopT forms, which are almost identical to the autocleaved version of NopTNGR234 (Khan et al., 2022). Remarkably, NopTUSDA257 and NopTHH103 were both able to cleave NFR5 in E. coli (Fig. 6A). However, both NopT proteins were unable to suppress cell death triggered by NFR1 and NFR5 expression in N. benthamiana leaves (Fig. 6B). Moreover, expression of NopTUSDA257 in NGR234ΔnopT had no effect on the infection ability of NGR234ΔnopT, as observed in inoculation tests with the L. japonicus line containing pNIN:GUS (Fig. 6C and 6D). These results suggest that NopTUSDA257 and NopTHH103 exhibit similar proteolytic activity to NopTNGR234, but target different host proteins to regulate rhizobial infection.

Revised version: Since NopT of S. fredii NGR234 (NopTNGR234) cleaves NFR5, we wondered whether homologs from other rhizobial species also possess such a proteolytic activity. Phylogenetic analysis showed that NopT proteins of Sinorhizobium and Bradyrhizobium are located in different clades (Fig. S15). B. diazoefficiens USDA110, a typical Bradyrhizobium strain used to study nodulation of soybeans, produces two NopT proteins (Fotiadis et al., 2012). In contrast to NopTNGR234, however, neither NopT1USDA110 nor NopT2USDA110 were able to cleave co-expressed NFR5 in E. coli cells (Fig. 6A). This finding suggests that NopT homologs from Bradyrhizbium species have lost the ability to cleave NFR5 and probably act on other, unknown host target proteins. The genome of various Sinorhizobium species (e.g., S. fredii USDA257 and S. fredii HH103) possess genes that encode truncated NopT forms, which are almost identical to the autocleaved version of NopTNGR234 (Khan et al., 2022). Remarkably, NopTUSDA257 and NopTHH103 were both able to cleave NFR5 in E. coli cells (Fig. 6A). However, NopTUSDA257 were unable to suppress cell death triggered by NFR1 and NFR5 expression in N. benthamiana leaves (Fig. 6B). Moreover, expression of NopTUSDA257 in NGR234ΔnopT had no effect on the infection ability of NGR234ΔnopT, as observed in inoculation tests with the L. japonicus line containing pNIN:GUS (Fig. 6C and 6D). These results suggest that NopTUSDA257 and NopTHH103 exhibit similar proteolytic activity to NopTNGR234, but target different host proteins to regulate rhizobial infection.

S. fredii NopT proteins cleave NFR5 and working model for NopT of NGR234

Earlier version: (A) NopT of S. fredii NGR234 and homologs from other rhizobial strains were co-expressed with NFR5CD in E. coli (NopT234, NopT of NGR234; NopT1110 and NopT2110, NopT proteins of B. diazoefficiens USDA110; NopT257, NopT of S. fredii USDA257; NopT103, NopT of S. fredii HH103). Immunoblot analysis indicated NFR5 cleavage by NopT proteins from S. fredii strains. (B) Expression of NopT257 in N. benthamiana could not inhibit the cell death triggered by co-expressed NFR1 and NFR5. (C) L. japonicus Gifu pNIN:GUS plants were inoculated with S. fredii NGR234, NGR234ΔnopT and NGR234ΔnopT expressing NopT257 (ΔnopT+NopT257). Roots were subjected to GUS staining at 7 dpi. Scale bar=2.5 mm. (D) Quantitative analysis of GUS-stained roots shown in panel C (n=19, Student’s t–test: P < 0.01). (E) A proposed model for NopT of NGR234 interacting with NFRs. NopT and NopTC (autocleaved NopT) proteolytically cleave NFR5 at the JM domain to release the intracellular domain of NFR5 (cleaved NFR5). NFR1 phosphorylates full-length NopT to block its proteinase activity. NopTC cannot be phosphorylated by NFR1.

Revised version: (A) NopT of S. fredii NGR234 and homologs from other rhizobial strains were co-expressed with NFR5CD in E. coli cells (NopT234, NopT of NGR234; NopT1110 and NopT2110, NopT proteins of B. diazoefficiens USDA110; NopT257, NopT of S. fredii USDA257; NopT103, NopT of S. fredii HH103). Immunoblot analysis indicated NFR5 cleavage by NopT proteins from S. fredii strains. (B) Expression of NopT257 in N. benthamiana could not inhibit the cell death triggered by co-expressed NFR1 and NFR5. (C) L. japonicus Gifu pNIN:GUS plants were inoculated with S. fredii NGR234, NGR234ΔnopT and NGR234ΔnopT expressing NopT257 (ΔnopT+NopT257). Roots were subjected to GUS staining at 7 dpi. Scale bar=2.5 mm. (D) Quantitative analysis of GUS-stained roots shown in panel C (n=19, Student’s t–test: P < 0.01). (E) A proposed model for NopT of NGR234 interacting with NFRs. NopT and NopTC (autocleaved NopT) proteolytically cleave NFR5 at the JM domain to release the intracellular domain of NFR5 (cleaved NFR5). NFR1 phosphorylates full-length NopT to block its proteinase activity. NopTC cannot be phosphorylated by NFR1.

Discussion

Earlier version: Plant-microbe interactions are complex, employing distinct “chemical weapons” or strategies to communicate mutually, thereby securing more resources for their survival (Jones and Dangl, 2006; Schubert et al., 2020). T3SS effectors of phytopathogenic bacteria are injected into plant cells to facilitate bacterial infection by targeting host proteins of the plant’s immune system, thereby suppressing PTI (Jones and Dangl, 2006; Tang et al., 2017). However, the role of rhizobial T3SS effectors in regulating the mutualistic symbiosis between legumes and rhizobia remains largely unknown. The aim of this study was to gain essential insights into the regulation mechanisms by which rhizobial effectors act directly on symbiotic signaling pathways to promote or dampen infection of host cells.

Revised version: Plant-microbe interactions are complex, employing distinct “chemical weapons” or strategies to communicate mutually, thereby securing more resources for their survival (Jones and Dangl, 2006; Schubert et al., 2020). T3SS effectors of phytopathogenic bacteria are injected into plant cells to facilitate bacterial infection by targeting host proteins of the plant immune system, thereby suppressing PTI (Jones and Dangl, 2006; Tang et al., 2017). However, the role of rhizobial T3SS effectors in regulating the mutualistic symbiosis between legumes and rhizobia remains largely unknown. The aim of this study was to gain essential insights into the regulation mechanisms by which rhizobial effectors act directly on symbiotic signaling pathways to promote or dampen infection of host cells.

Earlier version: Understanding how plants recognize different microbes as “friends” or “foes” is a central question in biology. The perception of microbial signals by plant receptor proteins is essential for triggering specific signaling pathways to resist invading pathogens or to establish symbiosis with beneficial microbes (Zipfel and Oldroyd, 2017). Direct targeting of these receptors by T3SS effectors represents an efficient strategy of phytopathogenic bacteria to suppress PTI and promote bacterial infection. In the well-studied Arabidopsis-Pseudomonas syringae interaction, T3SS effectors target flagellin receptor FLS2 and its coreceptor BAK1 (BRI1-[Brassinosteroid insensitive1]-associated receptor kinase 1) to suppress flagellin-induced PTI (Zipfel et al., 2004; Sun et al., 2013). For example, AvrPto targets FLS2 and HopB1 cleaves BAK1 (Xiang et al., 2008; Li et al., 2016). Likewise, AvrPtoB, an E3 ligase in P. syringae, targets FLS2 to promote its degradation and suppress PTI (Goehre et al., 2008). Despite the evolutionary importance of T3SS effectors in targeting host receptors, we were surprised to find that NopT, a rhizobial member of the C58 protease effector family, can directly target NFRs and manipulate the symbiotic signaling pathway in legumes. The interaction between NopT of S. fredii NGR234 and NFRs of L. japonicus was identified by using a unique screening system, in which co-expression of NFR1/NFR5 in N. benthamiana leaves leads to cell death. NopT, but not other effectors of NGR234, could suppress this cell death response. Consistent with these initial observations, inoculation of L. japonicus with the S. fredii NGR234ΔnopT mutant resulted in increased infection of L. japonicus roots. Moreover, over-expression of NopT in the roots was found to reduce rhizobial infection by M. loti. The interaction between NopT and NFR1/NFR5 was verified by BiFC, Split-LUC and co-IP experiments.

Revised version: Understanding how plants recognize different microbes as “friends” or “foes” is a central question in biology. The perception of microbial signals by plant receptor proteins is essential for triggering specific signaling pathways to resist invading pathogens or to establish symbiosis with beneficial microbes (Zipfel and Oldroyd, 2017). Direct targeting of these receptors by T3SS effectors represents an efficient strategy of phytopathogenic bacteria to suppress PTI and promote bacterial infection. In the well-studied Arabidopsis-Pseudomonas syringae interaction, T3SS effectors target plant flagellin receptor FLS2 and its coreceptor BAK1 (BRI1-[Brassinosteroid insensitive1]-associated receptor kinase 1) to suppress flagellin-induced PTI (Zipfel et al., 2004; Sun et al., 2013). For example, AvrPto targets FLS2 and HopB1 cleaves BAK1 (Xiang et al., 2008; Li et al., 2016). Likewise, AvrPtoB, an E3 ligase in P. syringae, targets FLS2 to promote its degradation and suppress PTI (Goehre et al., 2008). Despite the evolutionary importance of T3SS effectors in targeting host receptors, we were surprised to find that NopT, a rhizobial member of the C58 protease effector family, can directly target NFRs and manipulate the symbiotic signaling pathway in legumes. The interaction between NopT of S. fredii NGR234 and NFRs of L. japonicus was identified by using a unique screening system, in which co-expression of NFR1/NFR5 in N. benthamiana leaves leads to cell death. NopT, but not other effectors of NGR234, could suppress this cell death response. Consistent with these initial observations, inoculation of L. japonicus with the S. fredii NGR234ΔnopT mutant resulted in increased infection of L. japonicus roots. Moreover, over-expression of NopT in the roots was found to reduce rhizobial infection by M. loti. The interaction between NopT and NFR1/NFR5 was verified by BiFC, Split-LUC, and co-IP experiments.

Earlier version: Interestingly, using different assays, we found that NopT cleaves NFR5 at the JM domain. However, the efficiencies of NopT cleavage of NFR5 in L. japonicus and N. benthamiana were different (Fig. S14), which could be related to the uncontrollable activation of NFRs in N. benthamiana or phosphorylation of NopT by another kinase leading to decreased proteolytic activity. The modification of five conserved residues in the JM domain and the replacement of NFR5KC by NFPKC resulted in NFR5CD forms that were resistant to proteolysis by NopT in E. coli cells. These results suggest that cleavage of NFR5 depends on both the JM domain and the KC region. Although belonging to the same protease family, the NopT cleavage site in NFR5 was found to be polybasic as in YopT substrate (Shao et al., 2003a; Schmidt, 2011), whereas the cleavage site of AvrPphB substrate is canonical with seven adjacent amino acids involved (Kim et al., 2016). These differences reflect a broad variety of YopT-type effector proteases and their substrates in host cells. Our study shows, that NopT targets the NFR1/NFR5 complex thus reduces NF signaling and formation of nodule primordia in L. japonicus. NopT-NFR interactions may reflect a regulation mechanism to prevent rhizobial hyperinfection in host species other than L. japonicus.

Revised version: Interestingly, using different assays, we found that NopT cleaves NFR5 at the JM domain. However, the efficiencies of NopT cleavage of NFR5 in L. japonicus and N. benthamiana were slightly different (Fig. S16), which could be related to the uncontrollable activation of NFRs in N. benthamiana or phosphorylation of NopT by another kinase leading to decreased proteolytic activity. The modification of five conserved residues in the JM domain and the replacement of NFR5KC by NFPKC resulted in NFR5CD forms that were resistant to proteolysis by NopT in E. coli cells. These results suggest that cleavage of NFR5 depends on both the JM domain and the KC region. Although belonging to the same protease family, the NopT cleavage site in NFR5 was found to be polybasic as in the YopT substrate (Shao et al., 2003a; Schmidt, 2011), whereas the cleavage site of AvrPphB substrate is canonical with seven adjacent amino acids involved (Kim et al., 2016). These differences reflect a broad variety of YopT-type effector proteases and their substrates in host cells. Our study shows that NopT targets the NFR1/NFR5 complex thus reduces NF signaling and formation of nodule primordia in L. japonicus. NopT-NFR interactions may reflect a regulation mechanism to prevent rhizobial hyperinfection in host species other than L. japonicus.

Earlier version: Similar to the Pseudomonas effector AvrPphB, NopT of S. fredii NGR234 is able to cleave soybean PBS1-1 and expression of this effector in S. fredii USDA257 negatively influences symbiosis with certain soybean cultivars (Khan et al., 2022). Indeed, NopT of NGR234 acts as an Avr effector in some nonhost plants. Transient expression of NopT of strain NGR234 in Arabidopsis and tobacco triggers strong immune responses and cell death (Dai et al., 2008; Kimbrel et al., 2013; Khan et al., 2022; Fig. S15). In Arabidopsis, NopT-induced cell death was found to be dependent on the PBS1-1 kinase and the resistance protein RPS5 (Khan et al., 2022). Likewise, when constitutively expressing NopT and NopTC93S in L. japonicus, we only obtained stable transgenic lines expressing NopTC93S, indicating that the protease activity of NopT had a negative effect on plant development. In contrast, generation of hairy roots expressing NopT was possible. Such differences may be explained by different NopT substrates in roots and aerial parts of the plant. Indeed, our study shows that NopT not only cleaves NFR5 but is also able to proteolyze Arabidopsis AtLYK5 and L. japonicus LjLYS11. These two receptors possess chitin-binding affinities and trigger PTI responses (Cao et al., 2014; Gysel et al., 2021). NFR5, AtLYK5, and LjLYS11 are kinase-dead LYKs belonging to the same LYK subfamily. Thus, NopT appears not only to suppress NF signaling, but may also interfere with signal transduction pathways related to plant immunity. Depending on the host legume species, NopT could suppress PTI or induce ETI, thereby modulating rhizobial infection and nodule formation. Interactions between NopT and proteins related to the plant immune system may represent an important evolutionary driving force for host-specific nodulation and explain why the presence of NopT in NGR234 has a negative effect on symbiosis with L. japonicus but a positive one with other legumes.

Revised version: Similar to the Pseudomonas effector AvrPphB, NopT of S. fredii NGR234 is able to cleave soybean PBS1-1 and expression of this effector in S. fredii USDA257 negatively influences symbiosis with certain soybean cultivars (Khan et al., 2022). Indeed, NopT of NGR234 acts as an Avr effector in some nonhost plants. Transient expression of NopT of strain NGR234 in Arabidopsis and tobacco triggers strong immune responses and cell death (Dai et al., 2008; Kimbrel et al., 2013; Khan et al., 2022; Fig. S17). In Arabidopsis, NopT-induced cell death was found to be dependent on the PBS1-1 kinase and the resistance protein RPS5 (Khan et al., 2022). Likewise, when constitutively expressing NopT and NopTC93S in L. japonicus, we only obtained stable transgenic lines expressing NopTC93S, indicating that the protease activity of NopT had a negative effect on plant development. In contrast, generation of hairy roots expressing NopT was possible. Such differences may be explained by different NopT substrates in roots and aerial parts of the plant. Indeed, our study shows that NopT not only cleaves NFR5 but is also able to proteolyze Arabidopsis AtLYK5 and L. japonicus LjLYS11. These two receptors possess chitin-binding affinities and trigger PTI responses (Cao et al., 2014; Gysel et al., 2021). NFR5, AtLYK5, and LjLYS11 are kinase-dead LYKs belonging to the same LYK subfamily. Thus, NopT appears not only to suppress NF signaling, but may also interfere with signal transduction pathways related to plant immunity. Depending on the host legume species, NopT could suppress PTI or induce ETI, thereby modulating rhizobial infection and nodule formation. Interactions between NopT and proteins related to the plant immune system may represent an important evolutionary driving force for host-specific nodulation and explain why the presence of NopT in NGR234 has a negative effect on symbiosis with L. japonicus but a positive one with other legumes.

Earlier version: The regulation of kinase-dead pseudokinases in plants may be more complex compared to the well-studied control of protein kinase activities by phosphorylation. The cleavage of kinase-dead receptors may be an overlooked strategy to regulate their function in plant biology. Regulation of protein activities by cleavage is well studied in animals, e.g., for Notch signaling, where proteolysis of Notch molecules leads to downstream signaling (Kopan and Ilagan, 2009). Arabidopsis BAK1 undergoes proteolytic cleavage, which is important for both brassinosteroid signaling and induction of PTI responses (Zhou et al., 2019). The cytoplasmic domain of NFR5 is a pseudokinase that contributes to activation of NF signaling. Our work indicates that NopT proteolytically cleaves NFR5 and suppresses NF signaling. Such a fine-tuning process could be advantageous for the symbiosis in host plants other than L. japonicus, for example to prevent hyperinfection. However, the precise function of NFR5 cleavage by NopT and the fate of the released cytoplasmic domain remain to be clarified.

Revised version: The regulation of kinase-dead pseudokinases in plants may be more complex compared to the well-studied control of protein kinase activities by phosphorylation. The cleavage of receptor-like protein kinase may be an overlooked strategy to regulate their function in plant biology. Regulation of protein activities by cleavage is well studied in animals, e.g., for Notch signaling, where proteolysis of Notch molecules leads to downstream signaling (Kopan and Ilagan, 2009). Arabidopsis BAK1 undergoes proteolytic cleavage, which is important for both brassinosteroid signaling and induction of PTI responses (Zhou et al., 2019). Whereas, HopB1 cleaves immune co-receptor BAK1 at the kinase domain to inhibit plant defense (Li et al., 2016). NFR5 is a pseudokinase that is essentially required for the activation of NF signaling. However, the mechanisms of how the pseudokinase domain of NFR5 mediates symbiotic signaling transduction remains unclear. Our work indicates that NopT proteolytically cleaves NFR5 and suppresses NF signaling. Such a fine-tuning process could be advantageous for the symbiosis in host plants other than L. japonicus, for example to prevent hyperinfection. However, the precise function of NFR5 cleavage by NopT and the fate of the released cytoplasmic domain remain to be clarified.

Earlier version: Mutual regulation between T3SS effectors and host target proteins has been investigated for various plant-pathogen interactions. For example, the conserved Pseudomonas effector AvrPtoB can directly ubiquitinate the key components of plant immunity to promote bacterial virulence (Goehre et al., 2008). However, the activity of AvrPtoB in promoting virulence could either be enhanced or dampened by host proteins. AvrPtoB can be phosphorylated by SnRK2.8 to promote its virulence activity (Lei et al., 2020). On the other hand, the plant lectin receptor-like kinase LexRK-IX.2 and the Pto kinase phosphorylate AvrPtoB to inactivate its ubiquitin E3 ligase activity and undermine the effector’s ability to suppress PTI (Ntoukakis et al., 2009; Xu et al., 2020). A similar strategy is used by NFR1 to inhibit the protease activity of NopT via phosphorylation. NFR1 expressed in E. coli phosphorylates full-length NopT of S. fredii NGR234 at several residues. Three NopT forms with a phosphomimetic substitution (i.e., NopTT97D, NopTS119D, and NopTT257D) lacked autoproteolytic activity and the ability to cleave NFR5. The expression of these NopT variants in the NGR234ΔnopT mutant showed little effects in comparison with expression of wild-type NopT. These data suggest that L. japonicus utilizes NFR1 to phosphorylate NopT in order to dampen its catalytic activity and protect NFR5 from cleavage, which favors rhizobial infections. However, NopTΔN50, which is similar to autocleaved NopT, retained the ability to interact with NFR5 but not with NFR1. Moreover, full-length NopT, but not autocleaved NopT, migrated slower on gels when NFR1CD was co-expressed in E. coli, suggesting that autocleaved NopT may escape being phosphorylated and inactivated by NFR1. Future studies are required to explore the interactions between acylated NopT and NFR1. When expressed in N. benthamiana, GFP-tagged NopT forms lacking their autocleavage site (and thus non-acylated) were found to be localized within the cytoplasm, whereas NopT was targeted to the plasma membrane (Khan et al., 2022). Accordingly, NopT without acylation sites showed weaker interactions with NFRs in our study, suggesting that acylation of NopT, leading to plasma membrane location, promotes NopT-NFR interactions.

Revised version: Mutual regulation between T3SS effectors and host target proteins has been investigated for various plant-pathogen interactions. For example, the conserved Pseudomonas effector AvrPtoB can directly ubiquitinate the key components of plant immunity to promote bacterial virulence (Goehre et al., 2008). However, the activity of AvrPtoB in promoting virulence could either be enhanced or dampened by host proteins. AvrPtoB can be phosphorylated by SnRK2.8 to promote its virulence activity (Lei et al., 2020). On the other hand, the plant lectin receptor-like kinase LexRK-IX.2 and the Pto kinase phosphorylate AvrPtoB to inactivate its ubiquitin E3 ligase activity and undermine the effector’s ability to suppress PTI (Ntoukakis et al., 2009; Xu et al., 2020). A similar strategy is used by NFR1 to inhibit the protease activity of NopT via phosphorylation. NFR1 expressed in E. coli cells phosphorylates full-length NopT of S. fredii NGR234 at several residues. Three NopT forms with a phosphomimic substitution (i.e., NopTT97D, NopTS119D, and NopTT257D) lacked autoproteolytic activity and the ability to cleave NFR5. The expression of these NopT variants in the NGR234ΔnopT mutant showed little effects in comparison with expression of wild-type NopT. These data suggest that L. japonicus utilizes NFR1 to phosphorylate NopT in order to dampen its catalytic activity and protect NFR5 from cleavage, which favors rhizobial infections. However, NopTΔN50, which is similar to autocleaved NopT, retained the ability to interact with NFR5 but not with NFR1 (Fig. S3D). Moreover, full-length NopT, but not autocleaved NopT, migrated slower on gels when NFR1CD was co-expressed in E. coli cells, suggesting that autocleaved NopT may escape being phosphorylated and inactivated by NFR1. Future studies are required to explore the interactions between acylated NopT and NFR1. When expressed in N. benthamiana, GFP-tagged NopT forms lacking their autocleavage site (and thus non-acylated) were found to be localized within the cytoplasm, whereas NopT was targeted to the plasma membrane (Khan et al., 2022). Accordingly, NopT without acylation sites showed weaker interactions with NFRs in our study, suggesting that acylation of NopT, leading to plasma membrane location, promotes NopT-NFR interactions.

Earlier version: NopT homologs vary in different S. fredii strains. USDA257 and HH103 produce truncated NopT versions different from NGR234 due to a 19-bp deletion (Khan et al., 2022). Accordingly, NopT of these strains lacks autocleavage and acylation sites and possess a different N-terminal secretion signal sequence (Akeda and Galán, 2005). Nevertheless, NopTHH103 is a functional T3SS effector as mutant analysis indicated a symbiosis-promoting role in certain soybean cultivars (Li et al., 2023). Similar to NopTNGR234, NopTUSDA257 and NopTHH103 were able to cleave NFR5 in E. coli. Since NopTUSDA257 and NopTHH103 cannot be acylated, they likely accumulate in the cytoplasm of host legumes and are therefore less efficient in cleaving plasma membrane-bound NFR5 in comparison to acylated NopTNGR234. This could also explain why NopTUSDA257 was unable to suppress cell death triggered by NFR1/NFR5 in N. benthamiana leaves. It can be expected that NFR1 does not interact with and phosphorylate NopTUSDA257 or NopTHH103, as these effectors are similar to NopTΔN50 of NGR234, which cannot interact with NFR1. Consistent with these findings, expression of NopTUSDA257 in NGR234ΔnopT did not alter the infection phenotype of NGR234ΔnopT. It can be hypothesized that the host targets of these truncated NopT effectors might be different from those of NopTNGR234. Bradyrhizobium NopT effectors are also different from NopTNGR234, as NopT of B. diazoefficiens USDA110 was unable to cleave NFR5). NopT effectors may play a crucial role in Bradyrhizobium-legume interactions. A Bradyrhizobium sp. ORS3257 mutant deficient in NopT production induced only ineffective nodules on roots of Aeschynomene indica and most of the formed nodules were not infected (Teulet et al., 2019). Overall, these findings indicate versatile functions of NopT homologs in different strains and suggest that NopT controls different regulation processes in specific host cells.

Revised version: NopT homologs vary in different S. fredii strains. USDA257 and HH103 produce truncated NopT versions different from NGR234 due to a 19-bp deletion (Khan et al., 2022). Accordingly, NopT of these strains lacks autocleavage and acylation sites and possess a different N-terminal secretion signal sequence (Akeda and Galán, 2005). Nevertheless, NopTHH103 is a functional T3SS effector as mutant analysis indicated a symbiosis-promoting role in certain soybean cultivars (Li et al., 2023). Similar to NopTNGR234, NopTUSDA257 and NopTHH103 were able to cleave NFR5 in E. coli cells. Since NopTUSDA257 and NopTHH103 cannot be acylated, they likely accumulate in the cytoplasm of host legumes and are therefore less efficient in cleaving plasma membrane-bound NFR5 in comparison to acylated NopTNGR234. This could also explain why NopTUSDA257 was unable to suppress cell death triggered by NFR1/NFR5 in N. benthamiana leaves. It can be expected that NFR1 does not interact with and phosphorylate NopTUSDA257 or NopTHH103, as these effectors are similar to NopTΔN50 of NGR234, which cannot interact with NFR1. Consistent with these findings, expression of NopTUSDA257 in NGR234ΔnopT did not alter the infection phenotype of NGR234ΔnopT. It can be hypothesized that the host targets of these truncated NopT effectors might be different from those of NopTNGR234. Bradyrhizobium NopT effectors are also different from NopTNGR234, as NopT of B. diazoefficiens USDA110 was unable to cleave NFR5. NopT effectors may play a crucial role in Bradyrhizobium-legume interactions. A Bradyrhizobium sp. ORS3257 mutant deficient in NopT production induced only ineffective nodules on roots of Aeschynomene indica and most of the formed nodules were not infected (Teulet et al., 2019). Overall, these findings indicate versatile functions of NopT homologs in different strains and suggest that NopT controls different regulation processes in specific host cells.

Earlier version: We present here a model of mutual regulation, in which NopT proteolyzes NFR5 to suppress NF signaling, whereas protease activity of full-length NopT is suppressed by NFR1 via phosphorylation (Fig. 6E). NopT impairs the function of the NFR1/NFR5 receptor complex. Cleavage of NFR5 by NopT reduces its protein levels. Possible inhibitory effects of NFR5 cleavage products on NF signaling are unknown but cannot be excluded. Inactivation of NopT protease activity via phosphorylation by kinases such as NFR1 appears to be a countermeasure of the host to enable symbiotic signaling and rhizobial infection. Such feedback regulation could be the result of evolutionary pressure to produce ‘improved’ NFRs, which are resistant to cleavage by Not and have an increased capacity to phosphorylate NopT. On the other hand, suppression of NopT protease activity by NFRs may drive broad-host-range strains such as NGR234 to produce higher levels and structurally modified NFs.

Revised version: We present here a model of mutual regulation, in which NopT proteolyzes NFR5 to suppress NF signaling, whereas protease activity of full-length NopT is suppressed by NFR1 via phosphorylation (Fig. 6E). NopT impairs the function of the NFR1/NFR5 receptor complex. Cleavage of NFR5 by NopT reduces its protein levels. Possible inhibitory effects of NFR5 cleavage products on NF signaling are unknown but cannot be excluded. Inactivation of NopT protease activity via phosphorylation by kinases such as NFR1 appears to be a countermeasure of the host to enable symbiotic signaling and rhizobial infection. Such feedback regulation could be the result of evolutionary pressure to produce ‘improved’ NFRs, which are resistant to cleavage by NopT and have an increased capacity to phosphorylate NopT. On the other hand, suppression of NopT protease activity by NFRs may drive broad-host-range strains such as NGR234 to produce higher levels and structurally modified NFs.

Added in the revised version: The negative regulation of plant symbiotic signaling by NopT is detrimental to rhizobial infection, which appears paradoxical from an evolutionary perspective. Two possible scenarios will be discussed. One possibility is that NopT may have originated from its homologs, such as AvrPphB, and possesses a pathogenic feature by proteolyzing immune receptors to disrupt the plant immune pathway. However, due to sequence similarity, NFR5 is also proteolyzed by NopT. Expression of NopT leading to strong cell death in N. tabacum and Arabidopsis also suggests that NopT might be recogonized as Avr effector to trigger strong immunity. Another possibility is that the presence of NopT serves as a strategy for rhizobia to evade detection by plant immunity, which may be activated by rhizobial overinfection. The negative effects on symbiosis by cleaving NFR5 is then inhibited by NFR1 via a direct phosphorylation. In a parallel study, S. fredii NGR234 NopM was also identified to interact with and mediate the ubiquitination of NFR5, stabilizing NFR5 levels and thereby promoting rhizobial infection and nodulation (Wang et al., 2024). Therefore, both S. fredii NGR234 NopT and NopM target the same protein and appear to antagonistically regulate rhizobial symbiosis; however, the exact role of NopT in nodulation needs to be elucidated in future research.

Taken together, this study provides insights into a legume-rhizobium interaction, in which the bacterium deploys an effector protease to dampen symbiotic signaling, while the host plant counteracts by phosphorylating the effector, leading to its inactivation. Our findings highlight the function of a bacterial effector protease in regulating a symbiotic signaling pathway in legumes. This opens up the perspective of developing specific kinase-interacting proteases to reprogram and fine-tune cellular signaling processes in general.

Materials and Methods

Germination and growth of Lotus japonicus

Lotus japonicus (ecotype Gifu B-129) was provided by the Center for Carbohydrate Recognition and Signaling (https://lotus.au.dk/), and used as wild-type for nodulation assays. The Gifu pNIN:GUS transgenic line and nfr mutants (nfr1-1 and nfr5-3) were kindly provided by Dr. Jens Stougaard from the Aarhus University, Denmark (Madsen et al., 2003; Radutoiu et al. 2003; Heckmann et al., 2011). All L. japonicus seeds were treated with concentrated sulfuric acid for 10 min, followed by surface sterilization in 1% (w/v) NaClO for 8 min. After incubation at 4°C in the dark for 2 d, the seeds were placed on 0.8% (w/v) agar containing half-strength Murashige & Skoog (MS) medium for germination at 22°C for 2 d in the dark. Seedlings were then transferred to growth pots containing vermiculite supplied with half-strength B&D (Broughton & Dilworth) medium without nitrogen under long-day conditions (16-h light/8-h dark) at 22°C. Ten-day-old seedlings were inoculated with rhizobial cultures (OD600=0.02) for infection experiments.

Plasmid construction

Earlier version: All plasmids used in this study were generated using MultiF Seamless Assembly Mix (Abclonal, # RK21020). The plasmid pGWB514 was used as an original plasmid for construction of all binary vectors (Nakagawa et al., 2007). Briefly, different tags including HA tag, Myc tag, FLAG tag, Strep tag, nLUC, cLUC and GFP, were amplified using a forward primer containing a KpnI site and ligated into pGWB514 digested with XbaI and SacI using the seamless cloning method described above. The generated plasmids were named pG5XX-HA, pG5XX-Myc, pG5XX-FLAG, pG5XX-nLUC, pG5XX-cLUC, pG5XX-GFP, respectively. The coding sequences of L. japonicus NFR1 and NFR5 were cloned into pGWB514 and pG517 between XbaI/KpnI to generate NFR1-HA and NFR5-Myc fusion constructs under the control of the cauliflower mosaic virus 35S promoter. The DNA fragment containing pro35S:NFR5-Myc-NosT was then amplified and cloned into pG514-NFR1 at SbfI. The final plasmids contained pro35S:NFR1-HA-NosT and pro35S:NFR5-Myc-NosT. Rhizobial genes encoding effectors from S. fredii NGR234 were cloned into pG5XX-FLAG between XbaI/KpnI. To detect the interaction between NFR1/NFR5 and NopT in N. benthamiana, the full-length NFR1/NFR5, NopT and AtFLS2 coding sequences were cloned into pG5XX-Nluc, pG5XX-Nluc, pSPYCE(MR) and pSPYNE(R)173 (Waadt et al., 2008). For cleavage assays in N. benthamiana leaves, the NFR5 and NopT coding sequences were cloned into pG5XX-FLAG GFP and pG5XX-Strep, respectively. For cleavage assays in E. coli, the sequence encoding the cytoplasmic domain of NFR5 was cloned into pACYC duet Strep-HA (Han et al., 2017), and the sequences encoding the kinase domain of NFR5 fused to the HA tag and the JM domain of NFR5 fused to the GFP tag were cloned into pSUMO. The point mutations in NFR5CD and NopT were created by site-directed mutagenesis PCR using the templates pACYC duet Strep-NFR5CD-HA and pET28a NopT-FLAG. For NFR/NopT cleavage experiments with L. japonicus plants, NopT-FLAG or NopTC93S-FLAG was cloned into pUB-GFP between XbaI and KpnI, and 35S-NFR5-Myc-NosT was cloned into pUB-NopT-GFP at PstI. For Not phosphorylation experiments with E. coli, the sequence encoding the cytoplasmic domain of NFR1 fused to a Myc tag was cloned into pCDF duet. For complementation experiments with the NGR234ΔnopT mutant (Dai et al., 2008), a 1317-bp DNA fragment upstream of NopT and the coding sequence of a given NopT sequence was fused using overlap PCR and cloned into pHC60 between XhoI and KpnI. For complementation experiments with the nfr5-3 mutant, a 1316-bp NFR5 promoter DNA fragment was fused to NFR5-NFPKC or NFR55musing overlap PCR and cloned into pUB-Cherry between PstI and KpnI.

Revised version: All plasmids used in this study were generated using MultiF Seamless Assembly Mix (Abclonal, # RK21020). The plasmid pGWB514 was used as an original plasmid for construction of all binary vectors (Nakagawa et al., 2007). Briefly, different tags including HA tag, Myc tag, FLAG tag, Strep tag, NLuc, CLuc and GFP, were amplified using a forward primer containing a KpnI site and ligated into pGWB514 digested with XbaI and SacI using the seamless cloning method described above. The generated plasmids were named pG5XX-HA, pG5XX-Myc, pG5XX-FLAG, pG5XX-NLuc, pG5XX-CLuc, pG5XX-GFP, respectively. The coding sequences of L. japonicus NFR1 and NFR5 were cloned into pGWB514 and pG517 between XbaI/KpnI to generate NFR1-HA and NFR5-Myc fusion constructs under the control of the cauliflower mosaic virus 35S promoter. The DNA fragment containing pro35S:NFR5-Myc-NosT was then amplified and cloned into pG514-NFR1 at SbfI. The final plasmids contained pro35S:NFR1-HA-NosT and pro35S:NFR5-Myc-NosT. Rhizobial genes encoding effectors from S. fredii NGR234 were cloned into pG5XX-FLAG between XbaI/KpnI. To detect the interaction between NFR1/NFR5 and NopT in N. benthamiana, the full-length NFR1/NFR5, NopT and AtFLS2 coding sequences were cloned into pG5XX-NLuc, pG5XX-CLuc, pSPYCE (MR) and pSPYNE (R)173 (Waadt et al., 2008). For cleavage assays in N. benthamiana leaves, the NFR5 and NopT coding sequences were cloned into pG5XX-FLAG, pG5XX-GFP and pG5XX-Strep, respectively. For cleavage assays in E. coli cells, the sequence encoding the cytoplasmic domain of NFR5 was cloned into pACYC duet Strep-HA (Han et al., 2017), and the sequences encoding the kinase domain of NFR5 fused to the HA tag and the JM domain of NFR5 fused to the GFP tag were cloned into pSUMO. The point mutations in NFR5CD and NopT were created by site-directed mutagenesis PCR using the templates pACYC duet Strep-NFR5CD-HA and pET28a NopT-FLAG. For NFR/NopT cleavage experiments with L. japonicus plants, NopT-FLAG or NopTC93S-FLAG was cloned into pUB-GFP between XbaI and KpnI, and 35S-NFR5-Myc-NosT was cloned into pUB-NopT-GFP at PstI. For NopT phosphorylation experiments with E. coli, the sequence encoding the cytoplasmic domain of NFR1 fused to a Myc tag was cloned into pCDF duet. For complementation experiments with the NGR234ΔnopT mutant (Dai et al., 2008), a 1317-bp DNA fragment upstream of NopT and the coding sequence of a given NopT sequence was fused using overlap PCR and cloned into pHC60 between XhoI and KpnI. For complementation experiments with the nfr5-3 mutant, a 1316-bp NFR5 promoter DNA fragment was fused to NFR5-NFPKC or NFR55m using overlap PCR and cloned into pUB-Cherry between PstI and KpnI.

Transient gene expression in Nicotiana leaves

The plasmids of INF1, Avr3, BAX, and R3a for transient expression in Nicotiana benthamiana were kindly gifted by Dr. Juan Du from Huazhong Agricultural University. Agrobacterium tumefaciens strain EHA105 carrying various constructs were mixed with an Agrobacterium culture harboring the P19 suppressor in an infiltration buffer (10 mM MgCl2, 10 mM MES-KOH pH=5.8 and 200 μM acetosyringone). After incubation for 2 h at room temperature, Agrobacterium cultures were infiltrated into the leaves of 4-week-old N. benthamiana and N. tabacum plants. Leaves were harvested and analyzed two-days post inoculation. The cell death suppression experiments with N. benthamiana plants were performed as described previously (Wang et al., 2011).

Hairy root transformation

Earlier version: Hairy root transformation of L. japonicus plants was performed as described previously (Li et al., 2018). In brief, surface-sterilized seeds were germinated and grown on MS plates without sucrose (23°C in the dark for the first 3 days and 23°C at a 16-h light/8-h dark period for the following 2 days). The seeds were then cut at the middle of the hypocotyl and co-cultivated with A. rhizogenes LBA1334 carrying pUB-GFP, pUB-NopT-FLAG-GFP, pNFR5:NFR5-mCherry, pNFR5:NFR55m-mCherry and pNFR5:NFR5-NFPKC-mCherry, respectively (23°C in the dark for the first 3 days and 23°C at a 16-h light/8-h dark period for the following 2 days). The plants were then transferred onto solid B5 medium. After about 15 days, non-fluorescent hairy roots (lacking expression of GFP or mCherry) were removed. Where indicated, the plants were inoculated with rhizobia: S. fredii NGR234, NGR234ΔnopT or Mesorhizobium loti MAFF303099, and analyzed at indicated time points.

Revised version: Hairy root transformation of L. japonicus plants was performed as described previously (Li et al., 2018). In brief, surface-sterilized seeds were germinated and grown on MS plates without sucrose (23°C in the dark for the first 3 days and 23°C at a 16-h light/8-h dark period for the following 2 days). The seeds were then cut at the middle of the hypocotyl and co-cultivated with A. rhizogenes LBA1334 carrying pUB-GFP, pUB-NopT-FLAG-GFP, pNFR5:NFR5-mCherry, pNFR5:NFR55m-mCherry and pNFR5:NFR5-NFPKC-mCherry, respectively (23°C in the dark for the first 3 days and 23°C at a 16-h light/8-h dark period for the following 2 days). The plants were then transferred onto solid B5 medium. After about 15 days, non-fluorescent hairy roots (lacking expression of GFP or mCherry) were removed. The plants were inoculated with rhizobia: S. fredii NGR234, NGR234ΔnopT or Mesorhizobium loti MAFF303099, and analyzed at indicated time points.

Analysis of NFR5 cleavage by NopT and variants

Earlier version: In experiments with E. coli, Strep-NFR5CD-HA or Strep-NFR5CD-HA variants were co-expressed with NopT-FLAG or FLAG tagged NopT variant or AvrPphB-FLAG by using Duet vectors (Novagen; Han et al., 2017). Bacterial cultures were treated with 0.5 mM IPTG at 22°C for 15 h. After washing the bacterial cells with PBS buffer, the bacterial pellet was resuspended in extraction buffer consisting of Tris-HCl (pH 7.3), 150 mM NaCl, 5 mM EDTA, 0.5% (w/v) SDS and 0.5% (v/v) Triton X-100. After adding SDS loading buffer and boiling at 100°C for 5 min, the extracted proteins were subjected to immunoblot analysis using anti-HA-peroxidase (Sigma, clone 3F10) or anti-FLAG (Sigma, F1804) antibodies. An anti-groEL antibody (ABclonal, A0969) was used to confirm equal loading of proteins on gels.

Revised version: In experiments with E. coli, Strep-NFR5CD-HA or Strep-NFR5CD-HA variants were co-expressed with NopT-FLAG or FLAG tagged NopT variants or AvrPphB-FLAG by using Duet vectors (Novagen; Han et al., 2017). Bacterial cultures were treated with 0.5 mM IPTG at 22°C for 15 h. After washing the bacterial cells with PBS buffer, the bacterial pellet was resuspended in extraction buffer consisting of Tris-HCl (pH=7.3), 150 mM NaCl, 5 mM EDTA, 0.5% (w/v) SDS and 0.5% (v/v) Triton X-100. After adding SDS loading buffer and boiling at 100°C for 5 min, the extracted proteins were subjected to immunoblot analysis using anti-HA-peroxidase (Sigma, clone 3F10) or anti-FLAG (Sigma, F1804) antibodies. An anti-groEL antibody (ABclonal, A0969) was used to confirm equal loading of proteins on gels.

Earlier version: For in vitro cleavage assays, His-SUMO-NFR5JM-GFP and NopT-FLAG-His proteins were expressed in E. coli and purified. The reaction mixture contained 3 μg NopT-FLAG (or NopTC93S-FLAG) and 1 μg SUMO-NFR5JM-GFP in extraction buffer without SDS and EDTA. After incubation at 22°C for 7 h, 10 μL of a Ni-charged resin (GenScript, Cat. No. L00223) were added, and the cleavage products in the flow through were detected by immunoblot analysis using an anti-GFP antibody (ABclonal, AB_2770402).

Revised version: For in vitro cleavage assays, His-SUMO-NFR5JM-GFP and NopT-FLAG-His proteins were expressed in E. coli cells and purified. The reaction mixture contained 3 μg NopT-FLAG (or NopTC93S-FLAG) and 1 μg SUMO-NFR5JM-GFP in extraction buffer without SDS and EDTA. After incubation at 22°C for 7 h, 10 μL of a Ni-charged resin (GenScript, Cat. No. L00223) were added, and the cleavage products in the flow through were detected by immunoblot analysis using an anti-GFP antibody (ABclonal, AB_2770402).

Earlier version: To analyze cleavage of NFR5 in L. japonicus roots, hairy roots co-expressing NopT-FLAG or NopTC93S-FLAG with NFR5-Myc were harvested at about 15 days after induction of hairy roots without rhizobia inoculation and subjected to immunoblot analysis. Expressed proteins were extract by extraction buffer and immunodetected with anti-FLAG antibodies and anti-Myc (Bio-Legend, # 626808) antibodies.

Revised version: To analyze cleavage of NFR5 in L. japonicus roots, hairy roots co-expressing NopT-FLAG or NopTC93S-FLAG with NFR5-Myc were harvested at about 15 days after induction of hairy roots without rhizobia inoculation and subjected to immunoblot analysis. Expressed proteins were extract by extraction buffer and immunodetected with anti-FLAG antibody and anti-Myc (Bio-Legend, # 626808) antibody.

For NFR5/NopT cleavage experiments with N. benthamiana plants, leaves expressing proteins (NopT-FLAG or NopTC93S-FLAG with NFR5-GFP) were analyzed. Leaf discs were collected at 2 days post infiltration with agrobacteria and extracted in the extraction buffer (Tris-HCl pH=7.3, 150 mM NaCl, 5 mM EDTA, 0.5% (w/v) SDS and 0.5% (v/v) Triton X-100) and 2% protease inhibitors (Sigma, P9599). NFR5 was immunodetected with anti-Myc and anti-GFP (ABclonal, AB_2770402) antibodies. The NopT or NopTC93S were detected with anti-FLAG antibody. The actin proteins were detected with anti-Actin antibody (Abclonal, # AC009).

Phosphorylation assays

Earlier version: For NopT phosphorylation in E. coli, His-NFR1CD-Myc or His-NFR1CD–K351E-Myc was co-expressed with NopT-FLAG-His by using Duet vectors (Novagen). Protein expression was induced by 0.5 mM IPTG and cells were grown at 28°C for 18 h. After simultaneous purification of NFR1CD and NopT using Ni-charged resin, the proteins were incubated with CIAP (TaKaRa, 2250B) at 37°C for 3 h. The mobility shifts of NFR1CD and NopT were detected with anti-Myc (Bio-Legend, 626808) and anti-FLAG (Sigma, F1804) antibodies by immunoblot analysis.

Revised version: For NopT phosphorylation in E. coli cells, His-NFR1CD-Myc or His-NFR1CD–K351E-Myc was co-expressed with NopT-FLAG-His by using Duet vectors (Novagen). Protein expression was induced by 0.5 mM IPTG and cells were grown at 28°C for 18 h. After simultaneous purification of NFR1CD and NopT using Ni-charged resin, the proteins were incubated with CIAP (TaKaRa, 2250B) at 37°C for 3 h. The mobility shifts of NFR1CD and NopT were detected with anti-Myc (Bio-Legend, 626808) and anti-FLAG (Sigma, F1804) antibodies by immunoblot analysis.

Earlier version: In phosphorylation experiments with L. japonicus plants, NopTC93S-FLAG was expressed in hairy roots of the nfr1-1 mutant and wild-type plants, respectively. After removal of roots without fluorescence signals, plants were kept on half-strength B&D medium for 5 days. Roots were inoculated with Mesorhizobium loti MAFF303099 or treated with water. Total proteins were extracted using the extraction buffer (Tris-HCl (pH 7.3), 150 mM NaCl, 5 mM EDTA, 0.5% (w/v) SDS and 0.5% (v/v) Triton X-100) supplemented with 2% protease inhibitor cocktail (Sigma, P9599) and 2% phosphatase inhibitor cocktail (Yeasen, 20109-A). Proteins were then precipitated with TCA to remove various contaminants (e.g., EDTA and surfactant)s. Phosphorylation of the NopTC93S protein was analyzed by 50 μM Zn2+-Phos-tag SDS-PAGE gel (Kato & Sakamoto, 2019) and detected with the anti-FLAG antibody by immunoblot analysis. As loading control, actin proteins were immunodetected with an anti-Actin antibody.

Revised version: In phosphorylation experiments with L. japonicus plants, NopTC93S-FLAG was expressed in hairy roots of the nfr1-1 mutant and wild-type plants, respectively. After removal of roots without fluorescence signals, plants were kept on half-strength B&D medium for 5 days. Roots were inoculated with Mesorhizobium loti MAFF303099 or treated with water. Total proteins were extracted using the extraction buffer (Tris-HCl pH=7.3, 150 mM NaCl, 5 mM EDTA, 0.5% (w/v) SDS and 0.5% (v/v) Triton X-100) supplemented with 2% protease inhibitor cocktail (Sigma, P9599) and 2% phosphatase inhibitor cocktail (Yeasen, 20109-A). Proteins were then precipitated with TCA to remove various contaminants (e.g., EDTA and surfactants). Phosphorylation of the NopTC93S protein was analyzed by 50 μM Zn2+-Phos-tag SDS-PAGE gel (Kato & Sakamoto, 2019) and detected with the anti-FLAG antibody by immunoblot analysis. As loading control, actin proteins were immunodetected with an anti-Actin antibody.

BiFC and Split–LUC assays

In the BiFC assay with N. benthamiana plants, the eYFP fluorescence of expressed fusion proteins was recorded 2-3 days post infiltration with A. tumefaciens using a confocal microscope (Leica TCS SP8). For the Split–LUC assay, N. benthamiana leaves were collected at 2 days post infiltration and sprayed with a solution of 1 mM D-luciferin (Promega, E1603) and 0.02% (v/v) Triton X-100. After a dark adaptation of 5 min, bioluminescence images were acquired using a Tanon Bio–Imaging System (Tanon, Shanghai, 4600).

LC-MS analysis

Earlier version: For analysis of NopT phosphorylation sites, NopT-FLAG-His was co-expressed with His labelled NFR1CD or NFR1CD-K351E in E. coli. The proteins were purified using Ni-charged resin (GenScript, Cat. No. L00223). The bands containing NopTP and NopT (control) proteins, were cut out, destained by 50 mM triethylammonium bicarbonate (TEAB) solution (50% (v/v) acetonitrile in 50 mM TEAB), and washed with 100% acetonitrile until the gel turned white. The protein bands were then in-gel digested with sequencing-grade trypsin. LC-MS analysis was performed by Novogene Co., Ltd. (Beijing, China).

Revised version: For analysis of NopT phosphorylation sites, NopT-FLAG-His was co-expressed with His labelled NFR1CD or NFR1CD-K351E in E. coli cells. The proteins were purified using Ni-charged resin (GenScript, Cat. No. L00223). The bands containing NopTP and NopT (control) proteins, were cut out, destained by 50 mM triethylammonium bicarbonate (TEAB) solution (50% (v/v) acetonitrile in 50 mM TEAB), and washed with 100% acetonitrile until the gel turned white. The protein bands were then in-gel digested with sequencing-grade trypsin. LC-MS analysis was performed by Novogene Co., Ltd. (Beijing, China).

Earlier version: For identificaiton of the cleavage site in NFR5, NopT and His-SUMO-NFR5CD-GST, were co-expressed in E. coli and purified using Glutathione Resin (GenScript, # L00206). About 80 μg of the cleaved protein was collected for analysis. The protein bands were digested with trypsin,glu-C, chymotrypsin and pepsin. LC-MS analysis was performed by Oulu Biotechnology Co., Ltd. (Shanghai, China).

Revised version: For identificaiton of the cleavage site in NFR5, NopT and His-SUMO-NFR5CD-GST, were co-expressed in E. coli cells and purified using Glutathione Resin (GenScript, # L00206). About 80 μg of the cleaved protein was collected for analysis. The protein bands were digested with trypsin, glu-C, chymotrypsin and pepsin. LC-MS analysis was performed by Oulu Biotechnology Co., Ltd. (Shanghai, China).

Examination of rhizobial infection

Earlier version: The morphology of rhizobial infection foci and infection threads have been described previously (Rae et al., 2021). In brief, an infection focus is formed when the tip of a root hair curls around a single bacterium, forming an infection pocket. The number of infection foci and infection threads was determined using a Leica fluorescence microscope (DM2500). Whole roots were mounted on microscope slides and infection foci and infection thread counts were expressed on a per-root basis.

Revised version: The morphology of rhizobial infection foci and infection threads have been described previously (Rae et al., 2021). An infection focus is formed when the tip of a root hair curls around a single bacterium, forming an infection pocket. The infection thread is a tubular structure forming in the root hair, then extends to cortical cells. The number of infection foci and infection threads was determined using a Leica fluorescence microscope (DM2500). Whole roots were mounted on microscope slides and infection foci and infection thread counts were expressed on a per-root basis.

Earlier version: Histochemical GUS staining of L. japonicus Gifu B-129 or Gifu pNIN:GUS roots inoculated with different rhizobial strains was performed using the chromogenic substrate 5-bromo-4-chloro-3-indolyl-β-D-glucuronide as described (Heckmann et al., 2011). After evacuation of air for 20LJmin, root samples were incubated in GUS staining solution overnight (at 37°C in the dark). The GUS-stained tissues were observed and photographed using a Motic Swift M200D compound microscope (Gifu B-129 roots inoculated with rhizobial strains harboring pT7-GUS) or a Nikon SMZ18 stereo microscope (Gifu pNIN:GUS inoculated with unlabeled rhizobia).

Revised version: Histochemical GUS staining of L. japonicus Gifu B-129 or Gifu pNIN:GUS roots inoculated with different rhizobial strains was performed using the chromogenic substrate 5-bromo-4-chloro-3-indolyl-β-D-glucuronide as described (Heckmann et al., 2011). After evacuation of air for 20 min, root samples were incubated in GUS staining solution overnight (at 37°C in the dark). The GUS-stained tissues were observed and photographed using a Motic Swift M200D compound microscope (Gifu B-129 roots inoculated with rhizobial strains harboring pT7-GUS) or a Nikon SMZ18 stereo microscope (Gifu pNIN:GUS inoculated with unlabeled rhizobia).